# Problem Statement
Anova Insurance, a global health insurance company, seeks to optimize its insurance policy premium pricing based on the health status of applicants. Understanding an applicant's health condition is crucial for two key decisions:
- Determining eligibility for health insurance coverage.
- Deciding on premium rates, particularly if the applicant's health indicates higher risks.

Your objective is to Develop a predictive model that utilizes health data to classify individuals as 'healthy' or 'unhealthy'. This classification will assist in making informed decisions about insurance policy premium pricing.

# Dataset Overview
The dataset contains 10,000 rows and 20 columns, including both numerical and categorical variables. Some columns have missing values, especially for older individuals, reflecting the scenario where certain health records may not be up-to-date. Here is the data dictionary.

- Age: Represents the age of the individual. Negative values seem to be present, which might indicate data entry errors or a specific encoding used for certain age groups.

- BMI (Body Mass Index): A measure of body fat based on height and weight. Typically, a BMI between 18.5 and 24.9 is considered normal.

- Blood_Pressure: Represents systolic blood pressure. Normal blood pressure is usually around 120/80 mmHg.

- Cholesterol: This is the cholesterol level in mg/dL. Desirable levels are usually below 200 mg/dL.

- Glucose_Level: Indicates blood glucose levels. It might be fasting glucose levels, with normal levels usually ranging from 70 to 99 mg/dL.

- Heart_Rate: The number of heartbeats per minute. Normal resting heart rate for adults ranges from 60 to 100 beats per minute.

- Sleep_Hours: The average number of hours the individual sleeps per day.

- Exercise_Hours: The average number of hours the individual exercises per day. 

- Water_Intake: The average daily water intake in liters.

- Stress_Level: A numerical representation of stress level.

- Target: This is a binary outcome variable, with '1' indicating 'Unhealthy' and '0' indicating 'Healthy'.

- Smoking: A categorical variable indicating smoking status. Contains values - (0,1,2) which specify the regularity of smoking with 0 being no smoking and 2 being regular smmoking.

- Alcohol: A categorical variable indicating alcohol consumption status. Contains values - (0,1,2) which specify the regularity of alcohol consumption with 0 being no consumption quality and 2 being regular consumption.

- Diet: A categorical variable indcating the quality of dietary habits. Contains values - (0,1,2) which specify the quality of the habit with 0 being poor diet quality and 2 being good quality.

- MentalHealth: Possibly a measure of mental health status. Contains values - (0,1,2) which specify the severity of the mental health with 0 being fine and 2 being highly severe

- PhysicalActivity: A categorical variable indicating levels of physical activity. Contains values - (0,1,2) which specify the instensity of the medical history with 0 being no Physical Activity and 2 being regularly active.

- MedicalHistory: Indicates the presence of medical conditions or history. Contains values - (0,1,2) which specify the severity of the medical history with 0 being nothing and 2 being highly severe.

- Allergies: A categorical variable indicating allergy status. Contains values - (0,1,2) which specify the severity of the allergies with 0 being nothing and 2 being highly severe.

- Diet_Type: Categorical variable indicating the type of diet an individual follows. Contains values(Vegetarian, Non-Vegetarian, Vegan).

- Blood_Group: Indicates the blood group of the individual Contains values (A, B, AB, O).

It is clear from the above description that the predictor variable is the 'Target' column. This time we will be creating a decision tree model.

Let us begin with importing the necessary libraries. And reading the preprocessed data.

```python
# Necessary library imports
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split
from sklearn.preprocessing import StandardScaler
from sklearn.tree import DecisionTreeClassifier
from sklearn.metrics import f1_score
```

```python
# Load the dataset
df = pd.read_csv('Healthcare_Dataset_Preprocessed.csv')
df.head()
```

Output:
```
Age   BMI  Blood_Pressure  Cholesterol  Glucose_Level  Heart_Rate  \
0   2.0  26.0           111.0        198.0           99.0        72.0   
1   8.0  24.0           121.0        199.0          103.0        75.0   
2  81.0  27.0           147.0        203.0          100.0        74.0   
3  25.0  21.0           150.0        199.0          102.0        70.0   
4  24.0  26.0           146.0        202.0           99.0        76.0   

   Sleep_Hours  Exercise_Hours  Water_Intake  Stress_Level  ...  Diet  \
0          4.0             1.0           5.0           5.0  ...     1   
1          2.0             1.0           2.0           9.0  ...     1   
2         10.0            -0.0           5.0           1.0  ...     2   
3          7.0             3.0           3.0           3.0  ...     1   
4         10.0             2.0           5.0           1.0  ...     2   

   MentalHealth  PhysicalActivity  MedicalHistory  Allergies  Diet_Type_Vegan  \
0             2                 1            
```

Now that we have a fairly processed dataset, let's separate the feature variable and the target variable. Let the target variable be 'y' and independent variables be 'X'.

```python
df.shape
```

Output:
```
(9549, 23)
```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 9549 entries, 0 to 9548
Data columns (total 23 columns):
 #   Column                Non-Null Count  Dtype  
---  ------                --------------  -----  
 0   Age                   9549 non-null   float64
 1   BMI                   9549 non-null   float64
 2   Blood_Pressure        9549 non-null   float64
 3   Cholesterol           9549 non-null   float64
 4   Glucose_Level         9549 non-null   float64
 5   Heart_Rate            9549 non-null   float64
 6   Sleep_Hours           9549 non-null   float64
 7   Exercise_Hours        9549 non-null   float64
 8   Water_Intake          9549 non-null   float64
 9   Stress_Level          9549 non-null   float64
 10  Target                9549 non-null   int64  
 11  Smoking               9549 non-null   int64  
 12  Alcohol               9549 non-null   int64  
 13  Diet                  9549 non-null   int64  
 14  MentalHealth          9549 non-null   int64  
 15  PhysicalActivity     
```

```python
df.columns
```

Output:
```
Index(['Age', 'BMI', 'Blood_Pressure', 'Cholesterol', 'Glucose_Level',
       'Heart_Rate', 'Sleep_Hours', 'Exercise_Hours', 'Water_Intake',
       'Stress_Level', 'Target', 'Smoking', 'Alcohol', 'Diet', 'MentalHealth',
       'PhysicalActivity', 'MedicalHistory', 'Allergies', 'Diet_Type_Vegan',
       'Diet_Type_Vegetarian', 'Blood_Group_AB', 'Blood_Group_B',
       'Blood_Group_O'],
      dtype='object')
```

<!-- grade_id: cell-01f6852d92d132ed, points: 0, grade: false, solution: true -->
```python
# Separating the independent and target variables
# your code here
X = df[['Age', 'BMI', 'Blood_Pressure', 'Cholesterol', 'Glucose_Level',
       'Heart_Rate', 'Sleep_Hours', 'Exercise_Hours', 'Water_Intake',
       'Stress_Level', 'Smoking', 'Alcohol', 'Diet', 'MentalHealth',
       'PhysicalActivity', 'MedicalHistory', 'Allergies', 'Diet_Type_Vegan',
       'Diet_Type_Vegetarian', 'Blood_Group_AB', 'Blood_Group_B',
       'Blood_Group_O']]
y = df[["Target"]]
```

<!-- grade_id: cell-7ee923dab2344a92, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
assert X.shape == (9549, 22), 'Make sure you set the _Target_ variable in a separate variable named y and independent variables be X.'
```

<h3>Performing train-test split</h3>

Now, let's split the data in a 70:30 ratio. Use the random state = 42 and use the variable names- X_train, X_test, y_train, y_test

<!-- grade_id: cell-17f1ea18258eba3e, points: 0, grade: false, solution: true -->
```python
# Split the data into train and test
# Write your code below
# your code here
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)
```

<!-- grade_id: cell-2e00ecf881929dcd, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
assert X_train.shape == (6684, 22), 'Make sure to use the test size as 0.3, random_state = 42, and split the data correctly'
```

With that we are done with the train-test split. Let us have a quick look at the shape of the data.

```python
# A quick look at the shape of the datasets
X_train.shape, X_test.shape, y_train.shape, y_test.shape
```

Output:
```
((6684, 22), (2865, 22), (6684, 1), (2865, 1))
```

<h3>Building the model</h3>

Now, let us create a decision tree classifier instance. We will set the random state to be 42 and save the instance in a variable named "DT_model".

```python
# Creating a decision tree classifier object called DT_model
DT_model = DecisionTreeClassifier(random_state = 42)
```

Now fit the model - DT_model with X_train and y_train.

<!-- grade_id: cell-2a8dc57fa8d6b4f1, points: 0, grade: false, solution: true -->
```python
# Building the model using the training data
# your code here
DT_model.fit(X_train,y_train)
```

Output:
```
DecisionTreeClassifier(random_state=42)
```

Now that our model is built, let's perfom prediction on both the train and test data.
This would help us to compare our model performance on train and test data.

```python
# Perfoming prediction on both the train and test data
y_pred_train = DT_model.predict(X_train)
y_pred_test = DT_model.predict(X_test)
```

```python
y_pred_train
```

Output:
```
array([1, 0, 0, ..., 0, 1, 0], dtype=int64)
```

```python
y_pred_test
```

Output:
```
array([0, 1, 0, ..., 0, 0, 0], dtype=int64)
```

Now let is find the f1-score for both train and test data. Store them in a variable named f1_train and f1_test.

<h3>Evaluating the model</h3>

Now let us get the f1 scores for both train and test data. Store the f1_scores for train data in 'f1_train' and 'f1_test' variables respectively.

<!-- grade_id: cell-9b0dab3998231a4d, points: 0, grade: false, solution: true -->
```python
# Get the train and test score 
# your code here
f1_train = f1_score(y_train, y_pred_train)
f1_train
```

Output:
```
1.0
```

```python
f1_test = f1_score(y_test, y_pred_test)
f1_test
```

Output:
```
0.9003300330033003
```

<!-- grade_id: cell-f13e87ba4e8fcf55, points: 4, grade: true, solution: false -->
```python
assert f1_train == 1.0, 'Make sure you have followed the conditions stated to train the model and predict the output'
```

As you can see, the model performance on the train data is exceptional. We have got a 100% performance for F1 score on train data. But the score seems low on test data hence leading to an overfitting model.

<h3>Calculating feature importance</h3>

To possibly counter this issue, Anova Insurance wants to find to find the important features and get rid of the irrelevant ones. You must store the value of importance in a dataframe named 'imp'.

<!-- grade_id: cell-b18ee19cd508ecf0, points: 0, grade: false, solution: true -->
```python
#Retreiving the score of each feature
feature_imp = DT_model.feature_importances_

#  Sort the features by importance in descending order in a dataframe.

# your code here
imp = pd.DataFrame({'Col_names': X_train.columns, 'Importance': np.round(feature_imp, 2)})\
.sort_values(by='Importance', ascending=False)
imp

# Calculate the cumulative sum of the 'Importance' column and store it in a new column called 'cum_imp'
imp['cum_imp'] = imp.Importance.cumsum()
imp
```

Output:
```
Col_names  Importance  cum_imp
1                    BMI        0.22     0.22
3            Cholesterol        0.17     0.39
2         Blood_Pressure        0.12     0.51
4          Glucose_Level        0.09     0.60
9           Stress_Level        0.09     0.69
6            Sleep_Hours        0.07     0.76
8           Water_Intake        0.06     0.82
0                    Age        0.05     0.87
5             Heart_Rate        0.04     0.91
7         Exercise_Hours        0.02     0.93
14      PhysicalActivity        0.01     0.94
16             Allergies        0.01     0.95
15        MedicalHistory        0.01     0.96
11               Alcohol        0.01     0.97
13          MentalHealth        0.00     0.97
12                  Diet        0.00     0.97
10               Smoking        0.00     0.97
17       Diet_Type_Vegan        0.00     0.97
18  Diet_Type_Vegetarian        0.00     0.97
19        Blood_Group_AB        0.00     0.97
20         Blood_Group_B        0.00     0.97
21 
```

<!-- grade_id: cell-73283a0df7f81739, points: 4, grade: true, solution: false -->
```python
assert imp.shape == (22,3), 'Make sure you the feature importance is saved in a variable named imp and sorted in ascending order of importance'
```

Now let us drop the columns where the cummulative importance is greater than 0.95. Save these column names in a variable named drop_col.

<!-- grade_id: cell-3cefe433443638bf, points: 0, grade: false, solution: true -->
```python
#identifying the variables with low importance
# your code here
drop_col = imp[imp.cum_imp >  0.95]['Col_names'].to_list()
drop_col
```

Output:
```
['Allergies',
 'MedicalHistory',
 'Alcohol',
 'MentalHealth',
 'Diet',
 'Smoking',
 'Diet_Type_Vegan',
 'Diet_Type_Vegetarian',
 'Blood_Group_AB',
 'Blood_Group_B',
 'Blood_Group_O']
```

<!-- grade_id: cell-d2c5c973698bfcc6, points: 0, grade: false, solution: true -->
```python
# Dropping the low important columns from X_train and X_test
# your code here
# Dropping the columns from X_train
X_train.drop(columns = drop_col, axis = 1, inplace = True)

# Dropping the columns from X_test
X_test.drop(columns = drop_col, axis = 1, inplace = True)
```

<!-- grade_id: cell-2bbde36f4253259c, points: 4, grade: true, solution: false -->
```python
# ignore this cell
assert X_train.shape == (6684, 11), 'Make sure you have dropped the unimportant variables in train data'
assert X_test.shape == (2865, 11), 'Make sure you have dropped the unimportant variables in test data'
```

<h3>Making predictions using the model</h3>

Based on the irrelevant columns we removed, let us create a new decision tree model to see check whether there is an improvement in the performance.

```python
# Building the decision tree classifier object
DT_new_model = DecisionTreeClassifier(random_state = 42)
```

Now fit the model - DT_new_model with X_train and y_train.

<!-- grade_id: cell-2548421db2248389, points: 0, grade: false, solution: true -->
```python
# Building the decision tree classifier object
# your code here
DT_new_model.fit(X_train,y_train)
```

Output:
```
DecisionTreeClassifier(random_state=42)
```

Now that our model is built, let's perfom prediction on both the train and test data.

<!-- grade_id: cell-fe5f06cec9e039ec, points: 0, grade: false, solution: true -->
```python
# Building the model using the training data
# your code here
y_pred_train = DT_new_model.predict(X_train)
y_pred_test = DT_new_model.predict(X_test)
```

<h3>Evaluating the model</h3>

As the next step, print the f1-scores for both train and test data. Save the f1 score for train data in 'f1_train_new' and f1 score for test data in 'f1_test_new'.

<!-- grade_id: cell-5e6e53d0a830b30d, points: 0, grade: false, solution: true -->
```python
# your code here
f1_train_new = f1_score(y_train, y_pred_train)
f1_train_new
```

Output:
```
1.0
```

```python
f1_test_new = f1_score(y_test, y_pred_test)
f1_test_new
```

Output:
```
0.9093926319283105
```

<!-- grade_id: cell-0d90317d9e34729a, points: 4, grade: true, solution: false -->
```python
# ignore this cell
assert f1_train_new == 1.0, 'Make sure you have followed the conditions stated to train the model and predict the output'
```

We can see that there has only been a marginal improvement in the model. Another way to improve the model is pruning the tree and explore by identifying the right depth.

For this, let's look at the depth of the current model.

```python
# Storing the depth of our current model
depth = DT_model.get_depth()
depth
```

Output:
```
19
```

Next, we will create a list of values, starting from the default tree depth and reducing by 3 until reaching a depth of 1. This list will help us perform pruning using different values of max_depth. Save this list in a variable named - 'max_depth_list'

<!-- grade_id: cell-98ca66b198138125, points: 0, grade: false, solution: true -->
```python
# List of max_depth values 

# your code here
max_depth_list = list(range(depth,0,-3))
```

```python
max_depth_list
```

Output:
```
[19, 16, 13, 10, 7, 4, 1]
```

<!-- grade_id: cell-92a81ab0e5c8cab2, points: 1, grade: true, solution: false -->
```python
# ignore this cell
assert len(max_depth_list) == 7, 'Make sure you have followed the conditions stated to train the model and predict the output'
```

Next, we will train Decision Tree models with different max_depth values and then compute and display both the train and test F1 scores upto three decimal points for each model. 

<!-- grade_id: cell-292decf49f11ebb2, points: 0, grade: false, solution: true -->
```python
# Dictionary to store the train and test F1 scores

f1_train_scores = {}
f1_test_scores = {}

# Loop through max_depth values and train the models
for depth in max_depth_list:
    # Initialize the Decision Tree model with the current max_depth value
    DT_model_prune = DecisionTreeClassifier(max_depth=depth, random_state=42)
    
    # Train the model
    # your code here
    DT_model_prune.fit(X_train,y_train)

    # Make predictions on the train dataset
    y_train_pred = DT_model_prune.predict(X_train)

    # Make predictions on the test dataset
    y_test_pred = DT_model_prune.predict(X_test)
    
    # Store the train F1 score in the dictionary with the max_depth as the key
    f1_train_scores[depth] = f1_score(y_train, y_train_pred)
    
    # Store the test F1 score in the dictionary with the max_depth as the key
    f1_test_scores[depth] = f1_score(y_test, y_test_pred)

# Print the train and test F1 scores for each model
for depth in max_depth_list:
    print(f"max_depth = {depth}|\
    Train Score = {f1_train_scores[depth]:.3f} |\
    Test score = {f1_test_scores[depth]:.3f}")
    print('_'*65)
```

Output:
```
max_depth = 19|    Train Score = 0.999 |    Test score = 0.914
_________________________________________________________________
max_depth = 16|    Train Score = 0.994 |    Test score = 0.917
_________________________________________________________________
max_depth = 13|    Train Score = 0.980 |    Test score = 0.918
_________________________________________________________________
max_depth = 10|    Train Score = 0.951 |    Test score = 0.914
_________________________________________________________________
max_depth = 7|    Train Score = 0.911 |    Test score = 0.898
_________________________________________________________________
max_depth = 4|    Train Score = 0.839 |    Test score = 0.849
_________________________________________________________________
max_depth = 1|    Train Score = 0.754 |    Test score = 0.760
_________________________________________________________________
```

<!-- grade_id: cell-be427a332f9b4267, points: 4, grade: true, solution: false -->
```python
# ignore this cell
assert 19 in f1_train_scores and f1_train_scores[19] > 0.99, "Train score is diffenrent from the expected score for max_depth = 19. Make sure you have followed the instructions carefully."
assert 13 in f1_train_scores and f1_train_scores[13] > 0.97, "Train score is diffenrent from the expected score for max_depth = 15. Make sure you have followed the instructions carefully."
assert 7 in f1_train_scores and f1_train_scores[7] > 0.91, "Train score is diffenrent from the expected score for max_depth = 7. Make sure you have followed the instructions carefully."
assert 1 in f1_train_scores and f1_train_scores[1] > 0.75, "Train score is diffenrent from the expected score for max_depth = 3. Make sure you have followed the instructions carefully."

```

We can see that at max_depth = 7, the test and train scores are almost same while keeping them both near 0.9, which is a good score. That proves that max_depth = 7 is the optimum value of depth we should choose for our model while making predictions!