### Graded Lab 2

Hello ! Welcome to Graded Lab of Module 2.
Here we will be working on an exciting business problem with help of interesting datasets.
Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answers to different questions asked by company officials.

# Instructor update: Feb 4, 2024

```python
def data_process(data1,data2,data3):
    # your code here
    return data
```

```python
# importing libraries
import pandas as pd
import numpy as np

### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-21a002bcfcc3cec1, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    # your code here
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-e52e2a36b201dfa4, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure that you have converted order_date into a datetime format correctly.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

### Q1. Return the top three subcategories that yield the best percentage profit? Return output as list of sub-categories, list of percentage values (rounded upto 2 decimals).
Hint :
Sub-Category profit percentage will be calculated as sub-category profit / Total profit 

<!-- grade_id: cell-baeb2bd6b317d7c2, points: 0, grade: false, solution: true -->
```python
def top_3_sales(data):
    # your code here
    return sub_categories,perc_values
```

<!-- grade_id: cell-3826b4716b6578dd, points: 5, grade: true, solution: false -->
```python
assert len(top_3_sales(data=sales)[0])==3,"Please include list of top 3 sub-categories only"
assert type(top_3_sales(data=sales)[0])==list,"Output type should be list only."
```

<!-- grade_id: cell-a849b08051e5b706, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q2.	Which city has the highest total sales ?

<!-- grade_id: cell-d9f21a4f52e97f2d, points: 0, grade: false, solution: true -->
```python
def top_city_sales(data):
    # your code here
```

<!-- grade_id: cell-a2adc54841783950, points: 5, grade: true, solution: false -->
```python
assert type(top_city_sales(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-4e536adeb5512294, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.3. In year 2013 , which country has reported lowest profit ? 
Hint: Calculate order year from order_date column.

<!-- grade_id: cell-d953dc7d357bedac, points: 0, grade: false, solution: true -->
```python
def lowest_profit_country(data):
    # your code here
```

<!-- grade_id: cell-c56e95301db9434c, points: 5, grade: true, solution: false -->
```python
assert type(lowest_profit_country(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-9751a68aabf608d7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.4. For which market, we could observe 2nd highest total discount ? Return output as market name.

<!-- grade_id: cell-2dcfa33d66fab066, points: 0, grade: false, solution: true -->
```python
def second_highest_discount_market(data):
    # your code here
```

<!-- grade_id: cell-e89b4f5212d9af8c, points: 5, grade: true, solution: false -->
```python
assert type(second_highest_discount_market(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-e9bfdf6a222a9834, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q5.	Within the subcategory ‘Copiers’ , which product was sold the most (in terms of quantity) and how many units were sold? Return output as a tuple (product name, quantities sold).

<!-- grade_id: cell-2372c765dbb61495, points: 0, grade: false, solution: true -->
```python
def copier_sales(data):
    # your code here
```

<!-- grade_id: cell-d88992c7c676a693, points: 5, grade: true, solution: false -->
```python
assert type(copier_sales(data=sales)[0])==str,"Please check the data type of answer , product name should be string."
assert type(copier_sales(data=sales)[1])==np.int64,"Please check the data type of answer , quantity sold should be an integer."
```

<!-- grade_id: cell-7523c1f6679181a6, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q6.	In 2014, which customer (identify by name) contributed to the highest total profit and how much was it (Round it to 4 digits) ? Return output as a tuple (customer name, profit)

<!-- grade_id: cell-2f80a755e05bee76, points: 0, grade: false, solution: true -->
```python
def cust_prof(data): 
    
    ### Extract order year from order date. You can use dt.year functions
    # your code here

    data['order_year']=None 
    
    ### Filter data for yr 2014
    sales_yr= None
    
    
    #### Aggregate the profits wrto customer name , you can make use of .groupby() function in pandas
    cust_profit=None
     
    ### Round the profit column to 4 decimal numbers  ,You can make use of round() function for rounding off 
    cust_profit['profit']=None
    
    ## sort the dataframe with decreasing order of profits.
    None
    
    ### Store customer_name with highest profit in variables below. 
    customer_name=None
    customer_profit=None
    
    
    return customer_name,customer_profit
```

<!-- grade_id: cell-0659fb2c54ca2128, points: 10, grade: true, solution: false -->
```python
assert type(cust_prof(data=sales)[0])==str,"Please check the data type of answer , customer name should be string."
assert type(cust_prof(data=sales)[1])==np.float64,"Please check the data type of answer , profit recorded should be float."
```

<!-- grade_id: cell-1fa5045852195c41, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q7.	How much do the different categories of items contribute to total sales? 

**Return output dataframe consisting 2 columns 'category' (product category),'sales_perc' (sales % values).**

**Make sure to round off sales percentage values to 2 decimals.**

<!-- grade_id: cell-aeeb4124ab318d78, points: 0, grade: false, solution: true -->
```python
def cat_sales_contri(data):
    # your code here
```

<!-- grade_id: cell-a509c09c77a96597, points: 5, grade: true, solution: false -->
```python
assert type(cat_sales_contri(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be a dataframe."
assert cat_sales_contri(data=sales).shape==(3,2),"Please check the data shape, total row count should be equal to number of unique categories."
```

<!-- grade_id: cell-bbb943a227501894, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q8. Can we identify the customers who have not made any purchases in the last 12 months, so that we can send them some promotional material to encourage them to come back and shop with us? 

**Return the output as a data frame with three columns: 'customer_name','customer_id' and 'Total_sales' (total sales amount that customer has accrued). Make sure that you are returning the mentioned column names in same sequence & as same spelled.**

**Hint:-**

You might think of working on this logic

1. Calculate latest order date for every customer
2. Calculate yr back order date by offsetting latest date by 365 days (you can make use of pd.DateOffset(days=) function)
3. Check if order date is in between yr_back_order date and latest_order_date, create binary flag.
4. For every customer,you will get multiple order dates which are before yr old order date. Aggregate the flag at user level.
5. Check for a customer if all the flags indicating that order date are before 1 yr then mark that customer as one who hasnt made any transations in last yr.
6. Take out its name by joining with customer data and aggregate it for finding total sales. 

<!-- grade_id: cell-3fca9a4c82b93bf0, points: 0, grade: false, solution: true -->
```python
def cust_purc(data):
    # your code here
    
    ### Groupby customer_id & consider latest order_date, name this dataframe as purchase_date. 
    purchase_date=None
    purchase_date.rename(columns={'order_date':'latest_purchase_date'},inplace=True)
    
    #### Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days , make use of pd.DateOffset() function.
    purchase_date['yr_back_order_date']=None
    
    ### Offset the latest_purchase_date by 1 more day in order to avoid intersection with latest_order_date.
    purchase_date['latest_purchase_date']=None
    
    ### We have purchase-date data ready with us so now its time to merge it with original sales data
    data=data.merge(purchase_date,on=None,how=None)
    
    #### Check if order_date falls in between 'yr_back_order_date','latest_purchase_date' , if 'yes' flag it as 1 else 0. 
    data['purchase_flag']=None

    #### Calculate sum of all flag per customer_id by using groupby() function, call this flag total as '# purchases_in_last_yr'
    purchase_in_lst_yr=None
    purchase_in_lst_yr.rename(columns={'purchase_flag':'# purchases_in_last_yr'},inplace=True)
    
    
    #### Merge the purchase_in_lst_yr with sales data.
    data=data.merge(purchase_in_lst_yr,on=None,how=None)
    
    #### Select the customers who have 0 in '# purchases_in_last_yr' as this will be customers who havent purchased anything in last yr.
    sales_df=None
    
    ### Its time to do final aggregation for selected customers cosnider their name , id & aggregate it on sales column for getting total sales.
    sales_df_1=None
    
    return sales_df_1[['customer_name','customer_id','Total_sales']] ### You can replace sales_df_1 by your dataframe name.

```

<!-- grade_id: cell-af43b1cfc50161bf, points: 10, grade: true, solution: false -->
```python
assert type(cust_purc(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be dataframe."
assert cust_purc(data=sales).shape==(109,3),"Please check the data shape."
```

<!-- grade_id: cell-a75ce88e766af09e, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```