### Graded Lab 2

Hello ! Welcome to Graded Lab of Module 2.
Here we will be working on an exciting business problem with help of interesting datasets.
Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answers to different questions asked by company officials.

```python
# importing libraries
import pandas as pd
import numpy as np

### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id  \
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
3       4  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
4       5  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   

        product_id         category sub_category  \
0  OFF-LA-10002782  Office Supplies       Labels   
1  FUR-FU-10004015        Furniture  Furnishings   
2  FUR-BO-10002352        Furniture    Bookcases   
3  OFF-BI-10004428  Office Supplies      Binders   
4  OFF-AR-10004594  Office Supplies          Art   

                         product_name   sales  quantity  discount  profit  \
0  Hon File Folder Labels, Adjustable   13.08         3       0.0    4.56   
1                Tenex Clock, Durable  252.16    
```

```python
pd.set_option('display.max_columns',None)
#pd.set_option('display.max_rows',None)
pd.set_option('display.width', 1000)
pd.set_option('display.float_format','{:.4f}'.format)
```

```python
cust.head()
```

Output:
```
customer_id   customer_name           city             state    country  postal_code      segment market         region
0    SC-20575    Sonia Cooley    Mexico City  Distrito Federal     Mexico          NaN     Consumer  LATAM    LATAM-North
1    KW-16570  Kelly Williams  Dos Quebradas         Risaralda   Colombia          NaN     Consumer  LATAM    LATAM-South
2    DP-13000   Darren Powers      São Paulo         São Paulo     Brazil          NaN     Consumer  LATAM    LATAM-South
3    TB-21250    Tim Brockman        Managua           Managua  Nicaragua          NaN     Consumer  LATAM  LATAM-Central
4    PK-18910    Paul Knutson        Managua           Managua  Nicaragua          NaN  Home Office  LATAM  LATAM-Central
```

```python
vend.head()
```

Output:
```
vendor vendor_code
0          Accenco      VE_001
1  Ace Corporation      VE_002
2             Nero      VE_003
3         Rushcorp      VE_004
4    Low Tide Corp      VE_005
```

*-----------------------------------------------------------------------------*

```python
df = sales.merge(right=vend, how="left", on="vendor_code")
```

```python
df.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570  FUR-BO-10002352        Furniture    Bookcases        Ikea 3-Shelf Cabinet, Mobile 193.2800         2    0.0000 54.0800         9.6300         Medium     
```

```python
df2 = df.merge(right=cust, how="left", on="customer_id")
```

```python
df2
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id        product_id         category sub_category                               product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code                vendor     customer_name           city             state   country  postal_code      segment  market       region
0           1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   OFF-LA-10002782  Office Supplies       Labels         Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001               Accenco      Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN     Consumer   LATAM  LATAM-North
1           2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   FUR-FU-10004015        Furniture  Furnishings                       Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium    
```

```python
df2["order_date"] = pd.to_datetime(df2['order_date'], format='%d/%m/%Y')
```

```python
df2.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Dos Quebradas 
```

```python
df2.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 51290 entries, 0 to 51289
Data columns (total 26 columns):
 #   Column          Non-Null Count  Dtype         
---  ------          --------------  -----         
 0   row_id          51290 non-null  int64         
 1   order_id        51290 non-null  object        
 2   order_date      51290 non-null  datetime64[ns]
 3   ship_date       51290 non-null  object        
 4   ship_mode       49720 non-null  object        
 5   customer_id     51290 non-null  object        
 6   product_id      51290 non-null  object        
 7   category        51290 non-null  object        
 8   sub_category    51290 non-null  object        
 9   product_name    49993 non-null  object        
 10  sales           51290 non-null  float64       
 11  quantity        51290 non-null  int64         
 12  discount        51290 non-null  float64       
 13  profit          51290 non-null  float64       
 14  shipping_cost   51290 non-null  float64       
 15  or
```

*-----------------------------------------------------------------------------*

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-21a002bcfcc3cec1, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    # your code here
    df = data1.merge(right=data3, how="left", on="vendor_code")
    df2 = df.merge(right=data2, how="left", on="customer_id")
    df2["order_date"] = pd.to_datetime(df2['order_date'], format='%d/%m/%Y')
    return df2
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-e52e2a36b201dfa4, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure that you have converted order_date into a datetime format correctly.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

```python
sales.columns
```

Output:
```
Index(['row_id', 'order_id', 'order_date', 'ship_date', 'ship_mode', 'customer_id', 'product_id', 'category', 'sub_category', 'product_name', 'sales', 'quantity', 'discount', 'profit', 'shipping_cost', 'order_priority', 'vendor_code', 'vendor', 'customer_name', 'city', 'state', 'country', 'postal_code', 'segment', 'market', 'region'], dtype='object')
```

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Dos Quebradas 
```

```python
sales.shape
```

Output:
```
(51290, 26)
```

### Q1. Return the top three subcategories that yield the best percentage profit? Return output as list of sub-categories, list of percentage values (rounded upto 2 decimals).

*-----------------------------------------------------------------------------*

```python
sales.groupby("sub_category")["profit"].sum().nlargest(n=3)
```

Output:
```
sub_category
Copiers     258567.5482
Phones      216717.0058
Bookcases   161924.4195
Name: profit, dtype: float64
```

```python
sales.groupby("sub_category")["profit"].sum().nlargest(n=3).keys()
```

Output:
```
Index(['Copiers', 'Phones', 'Bookcases'], dtype='object', name='sub_category')
```

```python
sales["profit"].sum()
```

Output:
```
1467457.29128
```

```python
top3 = pd.DataFrame(sales.groupby("sub_category", as_index=False)["profit"].sum())
top3
```

Output:
```
sub_category      profit
0   Accessories 129626.3062
1    Appliances 141680.5894
2           Art  57953.9109
3       Binders  72449.8460
4     Bookcases 161924.4195
5        Chairs 140396.2675
6       Copiers 258567.5482
7     Envelopes  29601.1163
8     Fasteners  11525.4241
9   Furnishings  46967.4255
10       Labels  15010.5120
11     Machines  58867.8730
12        Paper  59207.6827
13       Phones 216717.0058
14      Storage 108461.4898
15     Supplies  22583.2631
16       Tables -64083.3887
```

```python
top3["percent"] = (top3["profit"] / sales["profit"].sum()) * 100
```

```python
top3
```

Output:
```
sub_category      profit  percent
0   Accessories 129626.3062   8.8334
1    Appliances 141680.5894   9.6548
2           Art  57953.9109   3.9493
3       Binders  72449.8460   4.9371
4     Bookcases 161924.4195  11.0344
5        Chairs 140396.2675   9.5673
6       Copiers 258567.5482  17.6201
7     Envelopes  29601.1163   2.0172
8     Fasteners  11525.4241   0.7854
9   Furnishings  46967.4255   3.2006
10       Labels  15010.5120   1.0229
11     Machines  58867.8730   4.0116
12        Paper  59207.6827   4.0347
13       Phones 216717.0058  14.7682
14      Storage 108461.4898   7.3911
15     Supplies  22583.2631   1.5389
16       Tables -64083.3887  -4.3670
```

```python
sub_categories = list(top3.sub_category)
sub_categories
```

Output:
```
['Accessories',
 'Appliances',
 'Art',
 'Binders',
 'Bookcases',
 'Chairs',
 'Copiers',
 'Envelopes',
 'Fasteners',
 'Furnishings',
 'Labels',
 'Machines',
 'Paper',
 'Phones',
 'Storage',
 'Supplies',
 'Tables']
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-baeb2bd6b317d7c2, points: 0, grade: false, solution: true -->
```python
def top_3_sales(data):
        
    # Group by sub_category and sum the profits for each sub_category
    sub_category_profit = data.groupby("sub_category")["profit"].sum().reset_index()
    
    # Sort the subcategories by profit in descending order and select the top 3
    top3 = sub_category_profit.nlargest(3, "profit")
    
    # Calculate the percentage profit for each subcategory
    total_profit = sub_category_profit["profit"].sum()
    top3["percent"] = (top3["profit"] / total_profit * 100).round(2)
    
    # Extract the subcategories and their corresponding percentage values
    sub_categories = top3["sub_category"].tolist()
    perc_values = top3["percent"].tolist()
    
    return sub_categories, perc_values
```

<!-- grade_id: cell-3826b4716b6578dd, points: 5, grade: true, solution: false -->
```python
assert len(top_3_sales(data=sales)[0])==3,"Please include list of top 3 sub-categories only"
assert type(top_3_sales(data=sales)[0])==list,"Output type should be list only."
```

<!-- grade_id: cell-a849b08051e5b706, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q2.	Which city has the highest sales ?

*-----------------------------------------------------------------------------*

```python
sales.groupby("city", as_index=False)["sales"].sum().sort_values(by="sales", ascending=False)
```

Output:
```
city       sales
552  Santo Domingo 498503.8404
368        Managua 411722.4227
396    Mexico City 396207.2421
593    Tegucigalpa 296864.1822
535   San Salvador 203888.4652
..             ...         ...
192           Edea    166.2990
505         Rybnik    131.8080
451      Panevezys    124.4160
340        Larache     52.0470
363      Maiduguri      7.1730

[659 rows x 2 columns]
```

```python
df1 = sales.groupby("city", as_index=False)["sales"].sum().sort_values(by="sales", ascending=False)
df1
```

Output:
```
city       sales
552  Santo Domingo 498503.8404
368        Managua 411722.4227
396    Mexico City 396207.2421
593    Tegucigalpa 296864.1822
535   San Salvador 203888.4652
..             ...         ...
192           Edea    166.2990
505         Rybnik    131.8080
451      Panevezys    124.4160
340        Larache     52.0470
363      Maiduguri      7.1730

[659 rows x 2 columns]
```

```python
df1.iloc[0:1]["city"]
```

Output:
```
552    Santo Domingo
Name: city, dtype: object
```

```python
df1.iloc[0:1]["city"].astype("str")
```

Output:
```
552    Santo Domingo
Name: city, dtype: object
```

```python
df1
```

Output:
```
city       sales
552  Santo Domingo 498503.8404
368        Managua 411722.4227
396    Mexico City 396207.2421
593    Tegucigalpa 296864.1822
535   San Salvador 203888.4652
..             ...         ...
192           Edea    166.2990
505         Rybnik    131.8080
451      Panevezys    124.4160
340        Larache     52.0470
363      Maiduguri      7.1730

[659 rows x 2 columns]
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-d9f21a4f52e97f2d, points: 0, grade: false, solution: true -->
```python
def top_city_sales(data):
    
    df1 = data.groupby("city", as_index=False)["sales"].sum().sort_values(by="sales", ascending=False)
    
    return df1['city'].values[0]
```

<!-- grade_id: cell-a2adc54841783950, points: 5, grade: true, solution: false -->
```python
assert type(top_city_sales(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-4e536adeb5512294, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q.3. In year 2013 , which country has reported lowest profit ? 
**(Calculate order year from order_date column.)**

*-----------------------------------------------------------------------------*

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Dos Quebradas 
```

```python
sales["year"] = sales["order_date"].dt.year
```

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Do
```

```python
year2013 = sales[sales["year"] == 2013]
year2013
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id        product_id         category sub_category                               product_name    sales  quantity  discount   profit  shipping_cost order_priority vendor_code                vendor    customer_name       city      state  country  postal_code    segment  market       region  year
6           7  MX-2013-134096 2013-09-27  01/10/2013  Standard Class    DP-13000   OFF-EN-10001375  Office Supplies    Envelopes  GlobeWeis Mailers, with clear poly window  56.1200         2    0.0000  21.3200         4.1100         Medium      VE_006           Venusystems    Darren Powers  São Paulo  São Paulo   Brazil          NaN   Consumer   LATAM  LATAM-South  2013
7           8  MX-2013-134096 2013-09-27  01/10/2013  Standard Class    DP-13000   TEC-MA-10004956       Technology     Machines                   Konica Card Printer, Red 344.6400         3    0.0000 165.4200        26.2000         Medium      VE_007          
```

```python
year2013.describe()
```

Output:
```
row_id      sales   quantity   discount     profit  shipping_cost  postal_code       year
count 13799.0000 13799.0000 13799.0000 13799.0000 13799.0000     13799.0000       0.0000 13799.0000
mean  25192.3444   246.8111     3.4884     0.1403    29.4902        26.4186          NaN  2013.0000
std   15084.8198   494.7194     2.2802     0.2101   190.3613        57.0822          NaN     0.0000
min       7.0000     0.8360     1.0000     0.0000 -6599.9780         0.0100          NaN  2013.0000
25%   11931.5000    30.8160     2.0000     0.0000     0.0000         2.6200          NaN  2013.0000
50%   25167.0000    85.2320     3.0000     0.0000     9.2400         7.8300          NaN  2013.0000
75%   38423.5000   249.7200     5.0000     0.2000    36.8520        24.9650          NaN  2013.0000
max   51280.0000 17499.9500    14.0000     0.8000  8399.9760       923.6300          NaN  2013.0000
```

```python
year2013["profit"].min()
```

Output:
```
-6599.978
```

```python
countrymin = year2013["country"][year2013["profit"] == year2013["profit"].min()]
countrymin
```

Output:
```
39068    Mexico
Name: country, dtype: object
```

```python
countrymin.values[0]
```

Output:
```
'Mexico'
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-d953dc7d357bedac, points: 0, grade: false, solution: true -->
```python
def lowest_profit_country(data):
    
    # Filter the data to include only records from the year 2013
    year2013_data = data[data["order_date"].dt.year == 2013]
    
    # Find the row with the lowest profit in the year 2013
    lowest_profit_row = year2013_data[year2013_data["profit"] == year2013_data["profit"].min()]
    
    # Extract the country from the row with the lowest profit
    country_with_lowest_profit = lowest_profit_row["country"].values[0]
    
    return country_with_lowest_profit
```

<!-- grade_id: cell-c56e95301db9434c, points: 5, grade: true, solution: false -->
```python
assert type(lowest_profit_country(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-9751a68aabf608d7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q.4. For which market segment we could observe 2nd highest discount ?

*-----------------------------------------------------------------------------*

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Do
```

```python
market_segment_discount = sales.groupby("market")["discount"].count()
market_segment_discount
```

Output:
```
market
APAC          7
Africa     4798
Canada      313
EMEA       4884
EU           23
LATAM     41265
Name: discount, dtype: int64
```

```python
# Group the data by "market" and calculate the total discount for each segment
market_segment_discount = sales.groupby("market")["discount"].sum().reset_index()
market_segment_discount
```

Output:
```
market  discount
0    APAC    4.1000
1  Africa  799.6000
2  Canada   50.7000
3    EMEA  851.1000
4      EU    3.3200
5   LATAM 5620.9080
```

```python
# Sort the segments by discount in descending order
sorted_segments = market_segment_discount.sort_values(by="discount", ascending=False)
sorted_segments
```

Output:
```
market  discount
5   LATAM 5620.9080
3    EMEA  851.1000
1  Africa  799.6000
2  Canada   50.7000
0    APAC    4.1000
4      EU    3.3200
```

```python
# Select the second row (assuming at least two different segments exist)
second_highest_discount_segment = sorted_segments.iloc[1]["market"]
second_highest_discount_segment
```

Output:
```
'EMEA'
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-2dcfa33d66fab066, points: 0, grade: false, solution: true -->
```python
def second_highest_discount_market(data):
    market_segment_discount = data.groupby("market")["discount"].sum().reset_index()
    sorted_segments = market_segment_discount.sort_values(by="discount", ascending=False)
    second_highest_discount_segment = sorted_segments.iloc[1]["market"]
    
    return second_highest_discount_segment
```

<!-- grade_id: cell-e89b4f5212d9af8c, points: 5, grade: true, solution: false -->
```python
assert type(second_highest_discount_market(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-e9bfdf6a222a9834, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q5.	Which product was sold the most (in terms of quantity) within the subcategory ‘Copiers’ and how many units were sold? Return output as a tuple (product name, quantities sold)

*-----------------------------------------------------------------------------*

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Do
```

```python
copier = sales[sales["sub_category"] == "Copiers"]
copier
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id        product_id    category sub_category                   product_name    sales  quantity  discount    profit  shipping_cost order_priority vendor_code            vendor     customer_name           city              state             country  postal_code      segment  market         region  year
25         26  MX-2014-120607 2014-06-18  25/06/2014  Standard Class    SP-20650   TEC-CO-10003541  Technology      Copiers         HP Copy Machine, Color 326.1863         2    0.0020   64.7063        16.7800         Medium      VE_011    Energy Systems  Stephanie Phelps       Coyoacán   Distrito Federal              Mexico          NaN    Corporate   LATAM    LATAM-North  2014
32         33  US-2013-167129 2013-05-24  31/05/2013  Standard Class    SC-20050   TEC-CO-10004904  Technology      Copiers     Sharp Fax Machine, Digital 941.7782         8    0.4020 -601.7418        80.2800         Medium      VE_008         F
```

```python
copier.isna().sum()
```

Output:
```
row_id               0
order_id             0
order_date           0
ship_date            0
ship_mode           82
customer_id          0
product_id           0
category             0
sub_category         0
product_name        61
sales                0
quantity             0
discount             0
profit               0
shipping_cost        0
order_priority      47
vendor_code          0
vendor               0
customer_name        0
city                 0
state                0
country              0
postal_code       2223
segment              0
market               0
region               0
year                 0
dtype: int64
```

```python
copier["product_name"].value_counts()
```

Output:
```
Hewlett Copy Machine, Color           44
Brother Copy Machine, Color           40
HP Copy Machine, Color                36
Sharp Copy Machine, Color             32
Canon Fax Machine, Digital            25
                                      ..
Canon PC1060 Personal Laser Copier     4
Canon PC1080F Personal Copier          4
Sharp AL-1530CS Digital Copier         4
Canon PC940 Copier                     3
Canon Image Class D660 Copier          2
Name: product_name, Length: 129, dtype: int64
```

```python
most_sold_product = str(copier["product_name"].value_counts().idxmax())
most_sold_product
```

Output:
```
'Hewlett Copy Machine, Color'
```

```python
quantity_sold = np.int64(copier["product_name"].value_counts().max())
quantity_sold
```

Output:
```
44
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-2372c765dbb61495, points: 0, grade: false, solution: true -->
```python
def copier_sales(data):
    copier = data[data["sub_category"] == "Copiers"]
    
    # Find the product with the maximum quantity sold within the "Copiers" subcategory
    most_sold_product = str(copier["product_name"].value_counts().idxmax())
    quantity_sold = np.int64(copier["product_name"].value_counts().max())
    
    return most_sold_product, quantity_sold
```

```python
copier_sales(data=sales)[0]
```

Output:
```
'Hewlett Copy Machine, Color'
```

```python
copier_sales(data=sales)[1]
```

Output:
```
44
```

<!-- grade_id: cell-d88992c7c676a693, points: 5, grade: true, solution: false -->
```python
assert type(copier_sales(data=sales)[0])==str,"Please check the data type of answer , product name should be string."
assert type(copier_sales(data=sales)[1])==np.int64,"Please check the data type of answer , quantity sold should be an integer."
```

<!-- grade_id: cell-7523c1f6679181a6, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q6.	In 2014, which customer (identify by name) contributed to the highest total profit and how much was it (Round it to 4 digits) ? Return output as a tuple (customer name, profit)

*-----------------------------------------------------------------------------*

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Do
```

```python
sales.dtypes
```

Output:
```
row_id                     int64
order_id                  object
order_date        datetime64[ns]
ship_date                 object
ship_mode                 object
customer_id               object
product_id                object
category                  object
sub_category              object
product_name              object
sales                    float64
quantity                   int64
discount                 float64
profit                   float64
shipping_cost            float64
order_priority            object
vendor_code               object
vendor                    object
customer_name             object
city                      object
state                     object
country                   object
postal_code              float64
segment                   object
market                    object
region                    object
year                       int64
dtype: object
```

```python
sales["year"] = sales["order_date"].dt.year
```

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Do
```

```python
year2014 = sales[sales["year"] == 2014]
```

```python
year2014.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount   profit  shipping_cost order_priority vendor_code         vendor customer_name         city             state    country  postal_code      segment market         region  year
0        1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000   4.5600         1.0300         Medium      VE_001        Accenco  Sonia Cooley  Mexico City  Distrito Federal     Mexico          NaN     Consumer  LATAM    LATAM-North  2014
13      14  MX-2014-121923 2014-09-24  01/10/2014  Standard Class    PK-18910  OFF-EN-10004100  Office Supplies    Envelopes        Kraft Mailers, Security-Tint  80.1000         3    0.0000  37.6200         8.8600            Low      VE_006    Venusystems  Paul Knutson   
```

```python
ans = pd.DataFrame(year2014.groupby("customer_name", as_index=False)["profit"].sum().nlargest(1, "profit"))
ans
```

Output:
```
customer_name    profit
622  Raymond Buch 7444.5343
```

```python
ans.iloc[:, 0]
```

Output:
```
622    Raymond Buch
Name: customer_name, dtype: object
```

```python
ans["customer_name"].item()
```

Output:
```
'Raymond Buch'
```

```python
ans["profit"].item()
```

Output:
```
7444.5343
```

```python
type(ans["profit"].item())
```

Output:
```
float
```

```python
ans.iloc[:, 1]
```

Output:
```
622   7444.5343
Name: profit, dtype: float64
```

```python
ans.iloc[:, 0] , ans.iloc[:, 1]
```

Output:
```
(622    Raymond Buch
 Name: customer_name, dtype: object,
 622   7444.5343
 Name: profit, dtype: float64)
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-2f80a755e05bee76, points: 0, grade: false, solution: true -->
```python
def cust_prof(data):      
    ### Extract order year from order date. You can use dt.year functions
    # your code here
    
    data['order_year'] = data["order_date"].dt.year
    
    ### Filter data for yr 2014
    sales_yr= data[data["order_year"] == 2014]
    
    
    #### Aggregate the profits wrto customer name , you can make use of .groupby() function in pandas
    cust_profit = sales_yr.groupby('customer_name')['profit'].sum().reset_index()
    #print(cust_profit) 
    ### Round the profit column to 4 decimal numbers  ,You can make use of round() function for rounding off 
    cust_profit['profit'] = cust_profit['profit'].round(4)
    #print(cust_profit['profit'])
    ## sort the dataframe with decreasing order of profits.
    highest_profit_customer = cust_profit.nlargest(1, 'profit')
    #print(highest_profit_customer)
    ### Store customer_name with highest profit in variables below. 
    customer_name = highest_profit_customer['customer_name'].values[0]
    customer_profit = highest_profit_customer['profit'].values[0]
    
    
    return customer_name,customer_profit
```

<!-- grade_id: cell-0659fb2c54ca2128, points: 10, grade: true, solution: false -->
```python
assert type(cust_prof(data=sales)[0])==str,"Please check the data type of answer , customer name should be string."
assert type(cust_prof(data=sales)[1])==np.float64,"Please check the data type of answer , profit recorded should be float."
```

<!-- grade_id: cell-1fa5045852195c41, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q7.	How much do the different categories of items contribute to total sales? 
**Return output dataframe consisting 2 columns 'category' (product category),'sales_perc' (sales % values).**

**Make sure to round off sales percentage values to 2 decimals.**

*-----------------------------------------------------------------------------*

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corpora
```

```python
# Group the data by 'category' and calculate the sum of sales for each category
category_sales = sales.groupby('category')['sales'].sum().reset_index()
category_sales
```

Output:
```
category        sales
0        Furniture 4110874.1863
1  Office Supplies 3787070.2255
2       Technology 4744557.4981
```

```python
# Calculate the total sales across all categories
total_sales = category_sales['sales'].sum()
total_sales
```

Output:
```
12642501.909880001
```

```python
# Calculate the sales percentage for each category and round to 2 decimals
category_sales['sales_perc'] = (category_sales['sales'] / total_sales * 100).round(2)
category_sales
```

Output:
```
category        sales  sales_perc
0        Furniture 4110874.1863     32.5200
1  Office Supplies 3787070.2255     29.9600
2       Technology 4744557.4981     37.5300
```

```python
category_sales.set_index(keys="category", inplace=True)
```

```python
category_sales
```

Output:
```
sales  sales_perc
category                                
Furniture       4110874.1863     32.5200
Office Supplies 3787070.2255     29.9600
Technology      4744557.4981     37.5300
```

*-----------------------------------------------------------------------------*

<!-- grade_id: cell-aeeb4124ab318d78, points: 0, grade: false, solution: true -->
```python
def cat_sales_contri(data):
    # Group the data by 'category' and calculate the sum of sales for each category
    category_sales = data.groupby('category')['sales'].sum().reset_index()
    
    # Calculate the total sales across all categories
    total_sales = category_sales['sales'].sum()
    
    # Calculate the sales percentage for each category and round to 2 decimals
    category_sales['sales_perc'] = (category_sales['sales'] / total_sales * 100).round(2)
    
    category_sales.set_index(keys="category", inplace=True)
    
    return category_sales
```

<!-- grade_id: cell-a509c09c77a96597, points: 5, grade: true, solution: false -->
```python
assert type(cat_sales_contri(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be a dataframe."
assert cat_sales_contri(data=sales).shape==(3,2),"Please check the data shape, total row count should be equal to number of unique categories."
```

<!-- grade_id: cell-bbb943a227501894, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

### Q8. Can we identify the customers who have not made any purchases in the last 12 months, so that we can send them some promotional material to encourage them to come back and shop with us? 

**Return the output as a data frame with three columns: 'customer_name','customer_id' and 'Total_sales' (total sales amount that customer has accrued). Make sure that you are returning the mentioned column names in same sequence & as same spelled.**

**Hint:-**

You might think of working on this logic

1. Calculate latest order date for every customer
2. Calculate yr back order date by offsetting latest date by 365 days (you can make use of pd.DateOffset(days=) function)
3. Check if order date is in between yr_back_order date and latest_order_date, create binary flag.
4. For every customer,you will get multiple order dates which are before yr old order date. Aggregate the flag at user level.
5. Check for a customer if all the flags indicating that order date are before 1 yr then mark that customer as one who hasnt made any transations in last yr.
6. Take out its name by joining with customer data and aggregate it for finding total sales. 

```python
sales.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 51290 entries, 0 to 51289
Data columns (total 28 columns):
 #   Column          Non-Null Count  Dtype         
---  ------          --------------  -----         
 0   row_id          51290 non-null  int64         
 1   order_id        51290 non-null  object        
 2   order_date      51290 non-null  datetime64[ns]
 3   ship_date       51290 non-null  object        
 4   ship_mode       49720 non-null  object        
 5   customer_id     51290 non-null  object        
 6   product_id      51290 non-null  object        
 7   category        51290 non-null  object        
 8   sub_category    51290 non-null  object        
 9   product_name    49993 non-null  object        
 10  sales           51290 non-null  float64       
 11  quantity        51290 non-null  int64         
 12  discount        51290 non-null  float64       
 13  profit          51290 non-null  float64       
 14  shipping_cost   51290 non-null  float64       
 15  or
```

```python
sales.shape
```

Output:
```
(51290, 28)
```

### Using CGPT

```python
# Group by customer_id and consider the latest order_date
purchase_date = sales.groupby('customer_id')['order_date'].max().reset_index()
purchase_date.rename(columns={'order_date': 'latest_purchase_date'}, inplace=True)
purchase_date
```

Output:
```
customer_id latest_purchase_date
0       AA-10315           2014-12-23
1       AA-10375           2014-12-25
2       AA-10480           2014-08-28
3       AA-10645           2014-12-03
4         AA-315           2014-12-29
...          ...                  ...
1585    YS-21880           2014-12-22
1586    ZC-11910           2014-06-14
1587    ZC-21910           2014-12-28
1588    ZD-11925           2014-12-28
1589    ZD-21925           2014-12-30

[1590 rows x 2 columns]
```

```python
# Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days
purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
purchase_date
```

Output:
```
customer_id latest_purchase_date yr_back_order_date
0       AA-10315           2014-12-23         2013-12-23
1       AA-10375           2014-12-25         2013-12-25
2       AA-10480           2014-08-28         2013-08-28
3       AA-10645           2014-12-03         2013-12-03
4         AA-315           2014-12-29         2013-12-29
...          ...                  ...                ...
1585    YS-21880           2014-12-22         2013-12-22
1586    ZC-11910           2014-06-14         2013-06-14
1587    ZC-21910           2014-12-28         2013-12-28
1588    ZD-11925           2014-12-28         2013-12-28
1589    ZD-21925           2014-12-30         2013-12-30

[1590 rows x 3 columns]
```

```python
# Offset the latest_purchase_date by 1 more day to avoid intersection with latest_order_date
purchase_date['latest_purchase_date'] =  purchase_date['latest_purchase_date'] + pd.DateOffset(days=1)
purchase_date
```

Output:
```
customer_id latest_purchase_date yr_back_order_date
0       AA-10315           2014-12-24         2013-12-23
1       AA-10375           2014-12-26         2013-12-25
2       AA-10480           2014-08-29         2013-08-28
3       AA-10645           2014-12-04         2013-12-03
4         AA-315           2014-12-30         2013-12-29
...          ...                  ...                ...
1585    YS-21880           2014-12-23         2013-12-22
1586    ZC-11910           2014-06-15         2013-06-14
1587    ZC-21910           2014-12-29         2013-12-28
1588    ZD-11925           2014-12-29         2013-12-28
1589    ZD-21925           2014-12-31         2013-12-30

[1590 rows x 3 columns]
```

```python
# Merge the purchase_date data with the original sales data
sales = sales.merge(purchase_date, on='customer_id', how='left')
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Tenex Clock, Durable 252.1600   
```

```python
# Check if the order_date falls between 'yr_back_order_date' and 'latest_purchase_date', if yes, flag it as 1, else 0
sales['purchase_flag'] = ((sales['order_date'] >= sales['yr_back_order_date']) & (sales['order_date'] <= sales['latest_purchase_date'])).astype(int)

sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                Te
```

```python
# Calculate the sum of all flags per customer_id
purchase_in_last_yr = sales.groupby('customer_id')['purchase_flag'].sum().reset_index()
purchase_in_last_yr.rename(columns={'purchase_flag': 'purchases_in_last_yr'}, inplace=True)
purchase_in_last_yr.head()
```

Output:
```
customer_id  purchases_in_last_yr
0    AA-10315                    18
1    AA-10375                    19
2    AA-10480                    11
3    AA-10645                    20
4      AA-315                     3
```

```python
(purchase_in_last_yr["purchases_in_last_yr"] == 0).sum()
```

Output:
```
0
```

```python
# Merge the purchase_in_last_yr with sales data
sales = sales.merge(purchase_in_last_yr, on='customer_id', how='left')

sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1                    17
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015    
```

```python
# Select the customers who have 0 in 'purchases_in_last_yr', as they haven't purchased anything in the last year
sales_df = sales[sales['purchase_flag'] == 0]
sales_df.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                               product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city      state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015        Furniture  Furnishings                       Tenex Clock, Durable 252.1600         8    0.0000 90.7200        13.4500         Medium      VE_002  Ace Corporation  Kelly Williams  Dos Quebradas  Risaralda  Colombia          NaN  Consumer  LATAM  LATAM-South  2012        2012           2014-12-10         2013-12-09              0                    22
2       3  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-BO-10002352    
```

```python
sales_df_1= sales_df.groupby(['customer_name', 'customer_id'])['sales'].sum().reset_index()
sales_df_1.head()
```

Output:
```
customer_name customer_id      sales
0   Aaron Bergman    AB-10015 13417.4433
1   Aaron Bergman       AB-15  1968.0600
2   Aaron Hawkins    AH-10030 11033.1078
3   Aaron Hawkins       AH-30  5036.5500
4  Aaron Smayling    AS-10045  4424.3600
```

```python
sales_df_1.rename(columns={'sales': 'Total_sales'}, inplace=True)
sales_df_1.head()
```

Output:
```
customer_name customer_id  Total_sales
0   Aaron Bergman    AB-10015   13417.4433
1   Aaron Bergman       AB-15    1968.0600
2   Aaron Hawkins    AH-10030   11033.1078
3   Aaron Hawkins       AH-30    5036.5500
4  Aaron Smayling    AS-10045    4424.3600
```

```python
sales_df_1[['customer_name','customer_id','Total_sales']]
```

Output:
```
customer_name customer_id  Total_sales
0          Aaron Bergman    AB-10015   13417.4433
1          Aaron Bergman       AB-15    1968.0600
2          Aaron Hawkins    AH-10030   11033.1078
3          Aaron Hawkins       AH-30    5036.5500
4         Aaron Smayling    AS-10045    4424.3600
...                  ...         ...          ...
1544      Yoseph Carroll    YC-11895      11.8500
1545      Yoseph Carroll    YC-21895   15180.1516
1546    Zuschuss Carroll    ZC-21910   23549.9803
1547  Zuschuss Donatelli    ZD-11925    2094.9660
1548  Zuschuss Donatelli    ZD-21925    7449.4055

[1549 rows x 3 columns]
```

<!-- grade_id: cell-3fca9a4c82b93bf0, points: 0, grade: false, solution: true -->
```python
def cust_purc(data):
    # your code here
    
    ### Groupby customer_id & consider latest order_date, name this dataframe as purchase_date. 
    purchase_date=data.groupby('customer_id')['order_date'].max().reset_index()
    purchase_date.rename(columns={'order_date':'latest_purchase_date'}, inplace=True)
    
    #### Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days , make use of pd.DateOffset() function.
    purchase_date['yr_back_order_date']= purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
    
    ### Offset the latest_purchase_date by 1 more day in order to avoid intersection with latest_order_date.
    purchase_date['latest_purchase_date'] += pd.DateOffset(days=1)
    
    ### We have purchase-date data ready with us so now its time to merge it with original sales data
    data=data.merge(purchase_date, on='customer_id', how='left')
    
    #### Check if order_date falls in between 'yr_back_order_date','latest_purchase_date' , if 'yes' flag it as 1 else 0. 
    data['purchase_flag'] = ((data['order_date'] >= data['yr_back_order_date']) & (data['order_date'] < data['latest_purchase_date'])).astype(int)

    #### Calculate sum of all flag per customer_id by using groupby() function, call this flag total as '# purchases_in_last_yr'
    purchase_in_last_yr=data.groupby('customer_id')['purchase_flag'].sum().reset_index()
    purchase_in_last_yr.rename(columns={'purchase_flag':'# purchases_in_last_yr'},inplace=True)
    
    
    #### Merge the purchase_in_lst_yr with sales data.
    data=data.merge(purchase_in_last_yr, on='customer_id', how='left')
    
    #### Select the customers who have 0 in '# purchases_in_last_yr' as this will be customers who havent purchased anything in last yr.
    sales_df= data[data['# purchases_in_last_yr'] == 0]
    
    ### Its time to do final aggregation for selected customers cosnider their name , id & aggregate it on sales column for getting total sales.
    sales_df_1 = sales_df.groupby(['customer_name', 'customer_id'])['sales'].sum().reset_index()
    sales_df_1.rename(columns={'sales': 'Total_sales'}, inplace=True)
    return sales_df_1[['customer_name','customer_id','Total_sales']] ### You can replace sales_df_1 by your dataframe name.

```

### Using Bing

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1                    17
1       2  MX-2012-155047 2012-10-15  20/10/2012  Standard Class    KW-16570  FUR-FU-10004015    
```

```python
# calculate latest order date for every customer
latest_order_date = sales.groupby('customer_id')['order_date'].max().reset_index()
latest_order_date
```

Output:
```
customer_id order_date
0       AA-10315 2014-12-23
1       AA-10375 2014-12-25
2       AA-10480 2014-08-28
3       AA-10645 2014-12-03
4         AA-315 2014-12-29
...          ...        ...
1585    YS-21880 2014-12-22
1586    ZC-11910 2014-06-14
1587    ZC-21910 2014-12-28
1588    ZD-11925 2014-12-28
1589    ZD-21925 2014-12-30

[1590 rows x 2 columns]
```

```python
# calculate year back order date by offsetting latest date by 365 days
year_back_order_date = latest_order_date.copy()
year_back_order_date['order_date'] = year_back_order_date['order_date'] - pd.DateOffset(years=1)
year_back_order_date
```

Output:
```
customer_id order_date
0       AA-10315 2013-12-23
1       AA-10375 2013-12-25
2       AA-10480 2013-08-28
3       AA-10645 2013-12-03
4         AA-315 2013-12-29
...          ...        ...
1585    YS-21880 2013-12-22
1586    ZC-11910 2013-06-14
1587    ZC-21910 2013-12-28
1588    ZD-11925 2013-12-28
1589    ZD-21925 2013-12-30

[1590 rows x 2 columns]
```

```python
# check if order date is in between year_back_order date and latest_order_date, create binary flag
sales = pd.merge(sales, latest_order_date, on = ['customer_id'], how='left', suffixes=('_original', '_latest'))
sales.head()
```

Output:
```
row_id        order_id order_date_original   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr order_date_latest
0       1  MX-2014-143658          2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1                    17        2014-12-04
1       2  MX-2012-155047          2012-10-
```

```python
sales = pd.merge(sales, year_back_order_date, on=['customer_id'], how='left', suffixes=('_original', '_back'))
sales.head()
```

Output:
```
row_id        order_id order_date_original   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr order_date_latest order_date
0       1  MX-2014-143658          2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1                    17        2014-12-04 2013-12-04
1       2  MX-2012-15
```

```python
sales.columns
```

Output:
```
Index(['row_id', 'order_id', 'order_date_original', 'ship_date', 'ship_mode', 'customer_id', 'product_id', 'category', 'sub_category', 'product_name', 'sales', 'quantity', 'discount', 'profit', 'shipping_cost', 'order_priority', 'vendor_code', 'vendor', 'customer_name', 'city', 'state', 'country', 'postal_code', 'segment', 'market', 'region', 'year', 'order_year', 'latest_purchase_date', 'yr_back_order_date', 'purchase_flag', 'purchases_in_last_yr', 'order_date_latest', 'order_date'], dtype='object')
```

```python
sales['flag'] = np.where((sales['order_date_original'] >= sales['order_date']) & (sales['order_date_original'] <= sales['order_date_latest']), 1, 0)
```

```python
sales.head()
```

Output:
```
row_id        order_id order_date_original   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name    sales  quantity  discount  profit  shipping_cost order_priority vendor_code           vendor   customer_name           city             state   country  postal_code   segment market       region  year  order_year latest_purchase_date yr_back_order_date  purchase_flag  purchases_in_last_yr order_date_latest order_date  flag
0       1  MX-2014-143658          2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001          Accenco    Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN  Consumer  LATAM  LATAM-North  2014        2014           2014-12-05         2013-12-04              1                    17        2014-12-04 2013-12-04     1
1       2
```

```python
# aggregate the flag at user level
sales_agg = sales.groupby('customer_id')['flag'].sum().reset_index()
sales_agg
```

Output:
```
customer_id  flag
0       AA-10315    18
1       AA-10375    19
2       AA-10480    11
3       AA-10645    20
4         AA-315     3
...          ...   ...
1585    YS-21880    21
1586    ZC-11910     1
1587    ZC-21910    27
1588    ZD-11925     8
1589    ZD-21925     6

[1590 rows x 2 columns]
```

```python
# check for a customer if all the flags indicating that order date are before 1 yr then mark that customer as one who hasn't made any transactions in last yr
sales_agg['no_purchase_last_year'] = np.where(sales_agg['flag'] == 0, 1, 0)
sales_agg
```

Output:
```
customer_id  flag  no_purchase_last_year
0       AA-10315    18                      0
1       AA-10375    19                      0
2       AA-10480    11                      0
3       AA-10645    20                      0
4         AA-315     3                      0
...          ...   ...                    ...
1585    YS-21880    21                      0
1586    ZC-11910     1                      0
1587    ZC-21910    27                      0
1588    ZD-11925     8                      0
1589    ZD-21925     6                      0

[1590 rows x 3 columns]
```

```python
# take out its name by joining with customer data and aggregate it for finding total sales
sales_agg = pd.merge(sales_agg, sales[['customer_id', 'customer_name', 'sales']].drop_duplicates(), on='customer_id', how='left')
sales_agg
```

Output:
```
customer_id  flag  no_purchase_last_year       customer_name    sales
0        AA-10315    18                      0          Alex Avila  61.3000
1        AA-10315    18                      0          Alex Avila 339.3600
2        AA-10315    18                      0          Alex Avila   9.0600
3        AA-10315    18                      0          Alex Avila  50.5000
4        AA-10315    18                      0          Alex Avila 172.6000
...           ...   ...                    ...                 ...      ...
51145    ZD-21925     6                      0  Zuschuss Donatelli 823.9600
51146    ZD-21925     6                      0  Zuschuss Donatelli  15.9840
51147    ZD-21925     6                      0  Zuschuss Donatelli  71.1200
51148    ZD-21925     6                      0  Zuschuss Donatelli 259.9600
51149    ZD-21925     6                      0  Zuschuss Donatelli  61.4400

[51150 rows x 5 columns]
```

```python
sales_agg
```

Output:
```
customer_id  flag  no_purchase_last_year       customer_name    sales
0        AA-10315    18                      0          Alex Avila  61.3000
1        AA-10315    18                      0          Alex Avila 339.3600
2        AA-10315    18                      0          Alex Avila   9.0600
3        AA-10315    18                      0          Alex Avila  50.5000
4        AA-10315    18                      0          Alex Avila 172.6000
...           ...   ...                    ...                 ...      ...
51145    ZD-21925     6                      0  Zuschuss Donatelli 823.9600
51146    ZD-21925     6                      0  Zuschuss Donatelli  15.9840
51147    ZD-21925     6                      0  Zuschuss Donatelli  71.1200
51148    ZD-21925     6                      0  Zuschuss Donatelli 259.9600
51149    ZD-21925     6                      0  Zuschuss Donatelli  61.4400

[51150 rows x 5 columns]
```

```python
sales_agg = sales_agg.groupby(['customer_id', 'customer_name'])['sales'].sum().reset_index()
sales_agg
```

Output:
```
customer_id       customer_name      sales
0       AA-10315          Alex Avila 13747.4130
1       AA-10375        Allen Armold  5884.1950
2       AA-10480        Andrew Allen 17643.7498
3       AA-10645       Anna Andreadi 15343.8907
4         AA-315          Alex Avila  2243.2560
...          ...                 ...        ...
1585    YS-21880       Yana Sorensen 17703.5860
1586    ZC-11910    Zuschuss Carroll     7.1730
1587    ZC-21910    Zuschuss Carroll 28272.8593
1588    ZD-11925  Zuschuss Donatelli  2951.2260
1589    ZD-21925  Zuschuss Donatelli  9479.3444

[1590 rows x 3 columns]
```

```python
sales_agg = sales_agg[sales_agg['no_purchase_last_year'] == 1][['customer_name', 'customer_id', 'sales']]
```

Output:
```
File C:\ProgramData\Anaconda3\lib\site-packages\pandas\core\frame.py:3807, in DataFrame.__getitem__(self, key)
   3805 if self.columns.nlevels > 1:
   3806     return self._getitem_multilevel(key)
-> 3807 indexer = self.columns.get_loc(key)
   3808 if is_integer(indexer):
   3809     indexer = [indexer]

File C:\ProgramData\Anaconda3\lib\site-packages\pandas\core\indexes\base.py:3804, in Index.get_loc(self, key, method, tolerance)
   3802     return self._engine.get_loc(casted_key)
   3803 except KeyError as err:
-> 3804     raise KeyError(key) from err
   3805 except TypeError:
   3806     # If we have a listlike key, _check_indexing_error will raise
   3807     #  InvalidIndexError. Otherwise we fall through and re-raise
   3808     #  the TypeError.
   3809     self._check_indexing_error(key)

KeyError: 'no_purchase_last_year'
```

<!-- grade_id: cell-3fca9a4c82b93bf0, points: 0, grade: false, solution: true -->
```python
def cust_purc(data):
    
    # calculate latest order date for every customer
    latest_order_date = data.groupby('customer_id')['order_date'].max().reset_index()

    # calculate year back order date by offsetting latest date by 365 days
    year_back_order_date = latest_order_date.copy()
    year_back_order_date['order_date'] = year_back_order_date['order_date'] - pd.DateOffset(years=1)
    
    # check if order date is in between year_back_order date and latest_order_date, create binary flag
    data = pd.merge(data, latest_order_date, on=['customer_id', 'order_date'], how='left')
    data = pd.merge(data, year_back_order_date, on=['customer_id'], how='left')
    data['flag'] = np.where((data['order_date'] >= data['order_date_y']) & (data['order_date'] <= data['order_date_x']), 1, 0)
```

### Using Bard

```python
# Groupby customer_id & consider latest order_date, name this dataframe as purchase_date
purchase_date = sales.groupby('customer_id')['order_date'].max().reset_index().rename(columns={'order_date': 'latest_purchase_date'})
purchase_date
```

```python
# Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days, make use of pd.DateOffset() function
purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
purchase_date
```

```python

```

```python

```

```python
def cust_purc(data):
    # Groupby customer_id & consider latest order_date, name this dataframe as purchase_date
    purchase_date = data.groupby('customer_id')['order_date'].max().reset_index().rename(columns={'order_date': 'latest_purchase_date'})

    # Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days, make use of pd.DateOffset() function
    purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)

    # Offset the latest_purchase_date by 1 more day in order to avoid intersection with latest_order_date
    purchase_date['latest_purchase_date'] = purchase_date['latest_purchase_date'] + pd.DateOffset(days=1)

    # Merge the purchase-date data with original sales data
    data = data.merge(purchase_date, on='customer_id', how='left')

    # Check if order_date falls in between 'yr_back_order_date','latest_purchase_date' , if 'yes' flag it as 1 else 0
    data['purchase_flag'] = data['order_date'].between(purchase_date['yr_back_order_date'], purchase_date['latest_purchase_date']).map({True: 1, False: 0})

    # Calculate sum of all flag per customer_id by using groupby() function, call this flag total as '# purchases_in_last_yr'
    purchase_in_lst_yr = data.groupby('customer_id')['purchase_flag'].sum().reset_index().rename(columns={'purchase_flag': '# purchases_in_last_yr'})

    # Merge the purchase_in_lst_yr with sales data
    data = data.merge(purchase_in_lst_yr, on='customer_id', how='left')

    # Select the customers who have 0 in '# purchases_in_last_yr' as this will be customers who havent purchased anything in last yr
    sales_df = data[data['# purchases_in_last_yr'] == 0]

    # Do final aggregation for selected customers cosnider their name , id & aggregate it on sales column for getting total sales
    sales_df_1 = sales_df.groupby(['customer_name', 'customer_id'])['sales'].sum().reset_index()

    return sales_df_1[['customer_name', 'customer_id', 'Total_sales']]
```

```python

```

<!-- grade_id: cell-af43b1cfc50161bf, points: 10, grade: true, solution: false -->
```python
assert type(cust_purc(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be dataframe."
assert cust_purc(data=sales).shape==(109,3),"Please check the data shape."
```

<!-- grade_id: cell-a75ce88e766af09e, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell
```

# Total Grade Obtained =  65/100 points earned

```python

```