### Graded Lab 2

Hello ! Welcome to Graded Lab of Module 2.
Here we will be working on an exciting business problem with help of interesting datasets.
Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answers to different questions asked by company officials.

# Instructor update: Feb 4, 2024 

```python
# importing libraries
import pandas as pd
import numpy as np

### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id  \
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
3       4  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
4       5  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   

        product_id         category sub_category  \
0  OFF-LA-10002782  Office Supplies       Labels   
1  FUR-FU-10004015        Furniture  Furnishings   
2  FUR-BO-10002352        Furniture    Bookcases   
3  OFF-BI-10004428  Office Supplies      Binders   
4  OFF-AR-10004594  Office Supplies          Art   

                         product_name   sales  quantity  discount  profit  \
0  Hon File Folder Labels, Adjustable   13.08         3       0.0    4.56   
1                Tenex Clock, Durable  252.16    
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-21a002bcfcc3cec1, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    
    # Merge data1 and data2 on 'customer_id'

    data = pd.merge(data1, data2, on='customer_id', how='left')

    # Merge the result with data3 on 'vendor_code'

    data = pd.merge(data, data3, on='vendor_code', how='left')

    # Convert 'order_date' into a datetime column

    data['order_date'] = pd.to_datetime(data['order_date'])

    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-e52e2a36b201dfa4, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure that you have converted order_date into a datetime format correctly.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

### Q1. Return the top three subcategories that yield the best percentage profit? Return output as list of sub-categories, list of percentage values (rounded upto 2 decimals).
Hint :
Sub-Category profit percentage will be calculated as sub-category profit / Total profit 

<!-- grade_id: cell-baeb2bd6b317d7c2, points: 0, grade: false, solution: true -->
```python
def top_3_sales(data):
    
    # Group by sub_category and sum the profits for each sub_category
    sub_category_profit = data.groupby("sub_category")["profit"].sum().reset_index()
    
    # Sort the subcategories by profit in descending order and select the top 3
    top3 = sub_category_profit.nlargest(3, "profit")
    
    # Calculate the percentage profit for each subcategory
    total_profit = sub_category_profit["profit"].sum()
    top3["percent"] = (top3["profit"] / total_profit * 100).round(2)
    
    # Extract the subcategories and their corresponding percentage values
    sub_categories = top3["sub_category"].tolist()
    perc_values = top3["percent"].tolist()
    
    return sub_categories, perc_values
```

<!-- grade_id: cell-3826b4716b6578dd, points: 5, grade: true, solution: false -->
```python
assert len(top_3_sales(data=sales)[0])==3,"Please include list of top 3 sub-categories only"
assert type(top_3_sales(data=sales)[0])==list,"Output type should be list only."
```

<!-- grade_id: cell-a849b08051e5b706, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
top_3_sales(data=sales)
```

Output:
```
(['Copiers', 'Phones', 'Bookcases'], [17.62, 14.77, 11.03])
```

```python
top_3_sales(data=sales)[0]   ### PASSED
```

Output:
```
['Copiers', 'Phones', 'Bookcases']
```

### Q2.	Which city has the highest total sales ?

<!-- grade_id: cell-d9f21a4f52e97f2d, points: 0, grade: false, solution: true -->
```python
def top_city_sales(data):
    
    # Group by city and calculate the total sales for each city
    city_total_sales = data.groupby('city')['sales'].sum()
    
    # Find the city with the highest sales
    top_city = city_total_sales.idxmax()

    return top_city
```

<!-- grade_id: cell-a2adc54841783950, points: 5, grade: true, solution: false -->
```python
assert type(top_city_sales(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-4e536adeb5512294, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
top_city_sales(sales) ### PASSED
```

Output:
```
'Santo Domingo'
```

### Q.3. In year 2013 , which country has reported lowest profit ? 
Hint: Calculate order year from order_date column.

<!-- grade_id: cell-d953dc7d357bedac, points: 0, grade: false, solution: true -->
```python
def lowest_profit_country(data):
    
    # Extract the year from the order_date column
    data['order_year'] = pd.to_datetime(data['order_date']).dt.year

    # Filter data for the year 2013
    data_2013 = data[data['order_year'] == 2013]
    
    # Group by country and calculate the total profit for each country
    country_total_profit = data_2013.groupby('country')['profit'].sum()

    # Find the country with the lowest profit
    lowest_profit_country = country_total_profit.idxmin()

    return lowest_profit_country
```

<!-- grade_id: cell-c56e95301db9434c, points: 5, grade: true, solution: false -->
```python
assert type(lowest_profit_country(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-9751a68aabf608d7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
lowest_profit_country(data=sales)  ### PASSED
```

Output:
```
'Togo'
```

### Q.4. For which market, we could observe 2nd highest total discount ? Return output as market name.

<!-- grade_id: cell-2dcfa33d66fab066, points: 0, grade: false, solution: true -->
```python
def second_highest_discount_market(data):
    
    # Group by market segment and calculate the mean discount for each segment
    segment_mean_discount = data.groupby('market')['discount'].mean()

    # Sort the market segments based on mean discount in descending order
    sorted_segments = segment_mean_discount.sort_values(ascending=False)

    # Select the market segment with the second-highest discount
    second_highest_discount_market = sorted_segments.index[1]

    return second_highest_discount_market
```

<!-- grade_id: cell-e89b4f5212d9af8c, points: 5, grade: true, solution: false -->
```python
assert type(second_highest_discount_market(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-e9bfdf6a222a9834, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
second_highest_discount_market(data=sales)  ### PASSED
```

Output:
```
'EMEA'
```

### Q5.	Within the subcategory ‘Copiers’ , which product was sold the most (in terms of quantity) and how many units were sold? Return output as a tuple (product name, quantities sold).

<!-- grade_id: cell-2372c765dbb61495, points: 0, grade: false, solution: true -->
```python
def copier_sales(data):
    
    # Filter data for the subcategory 'Copiers'

    copiers = data[data['sub_category'] == 'Copiers']

    # Group by product_name and calculate total quantity

    grouped = copiers.groupby('product_name').agg({'quantity': 'sum'})

    # Sort by quantity in descending order and take the top 1

    top_product = grouped.sort_values('quantity', ascending=False).head(1)

    # Return the name of the product and the quantity sold

    product = top_product.index[0]

    quantity_sold = top_product['quantity'].values[0]

    return product, quantity_sold
```

<!-- grade_id: cell-d88992c7c676a693, points: 5, grade: true, solution: false -->
```python
assert type(copier_sales(data=sales)[0])==str,"Please check the data type of answer , product name should be string."
assert type(copier_sales(data=sales)[1])==np.int64,"Please check the data type of answer , quantity sold should be an integer."
```

<!-- grade_id: cell-7523c1f6679181a6, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
copier_sales(data=sales)
```

Output:
```
('Hewlett Copy Machine, Color', 146)
```

```python
copier_sales(data=sales)[0]
```

Output:
```
'Hewlett Copy Machine, Color'
```

```python
copier_sales(data=sales)[1]  ### PASSED
```

Output:
```
146
```

### Q6.	In 2014, which customer (identify by name) contributed to the highest total profit and how much was it (Round it to 4 digits) ? Return output as a tuple (customer name, profit)

<!-- grade_id: cell-2f80a755e05bee76, points: 0, grade: false, solution: true -->
```python
def cust_prof(data): 
    
    ### Extract order year from order date. You can use dt.year functions
   

    data['order_year'] = pd.to_datetime(data['order_date']).dt.year
    
    ### Filter data for yr 2014
    sales_yr= data[data['order_year'] == 2014]
    
    
    #### Aggregate the profits wrto customer name , you can make use of .groupby() function in pandas
    cust_profit = sales_yr.groupby('customer_name')['profit'].sum().reset_index()
     
    ### Round the profit column to 4 decimal numbers  ,You can make use of round() function for rounding off 
    cust_profit['profit'] = round(cust_profit['profit'], 4)
    
    ## sort the dataframe with decreasing order of profits.
    cust_profit = cust_profit.sort_values(by='profit', ascending=False)
    
    ### Store customer_name with highest profit in variables below. 
    customer_name = cust_profit.iloc[0]['customer_name']
    customer_profit = cust_profit.iloc[0]['profit']
    
    
    return customer_name,customer_profit
```

<!-- grade_id: cell-0659fb2c54ca2128, points: 10, grade: true, solution: false -->
```python
assert type(cust_prof(data=sales)[0])==str,"Please check the data type of answer , customer name should be string."
assert type(cust_prof(data=sales)[1])==np.float64,"Please check the data type of answer , profit recorded should be float."
```

<!-- grade_id: cell-1fa5045852195c41, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
cust_prof(data=sales)
```

Output:
```
('Raymond Buch', 7444.5343)
```

```python
cust_prof(data=sales)[0]
```

Output:
```
'Raymond Buch'
```

```python
cust_prof(data=sales)[1]  ### PASSED
```

Output:
```
7444.5343
```

### Q7.	How much do the different categories of items contribute to total sales? 

**Return output dataframe consisting 2 columns 'category' (product category),'sales_perc' (sales % values).**

**Make sure to round off sales percentage values to 2 decimals.**

<!-- grade_id: cell-aeeb4124ab318d78, points: 0, grade: false, solution: true -->
```python
def cat_sales_contri(data):
    
    # Calculate total sales
    total_sales = data['sales'].sum()
    
    # Group by category and calculate the percentage of total sales
    cat_sales = data.groupby('category')['sales'].sum().reset_index()
    cat_sales['sales_perc'] = (cat_sales['sales'] / total_sales) * 100
    
    # Round off sales percentage values to 2 decimals
    cat_sales['sales_perc'] = round(cat_sales['sales_perc'], 2)
    
    # Select required columns
    cat_sales = cat_sales[['category', 'sales_perc']]
    
    return cat_sales
```

<!-- grade_id: cell-a509c09c77a96597, points: 5, grade: true, solution: false -->
```python
assert type(cat_sales_contri(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be a dataframe."
assert cat_sales_contri(data=sales).shape==(3,2),"Please check the data shape, total row count should be equal to number of unique categories."
```

<!-- grade_id: cell-bbb943a227501894, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
cat_sales_contri(data=sales)  
```

Output:
```
category  sales_perc
0        Furniture       32.52
1  Office Supplies       29.96
2       Technology       37.53
```

```python
type(cat_sales_contri(data=sales))
```

Output:
```
pandas.core.frame.DataFrame
```

```python
cat_sales_contri(data=sales).shape  ### PASSED
```

Output:
```
(3, 2)
```

### Q8. Can we identify the customers who have not made any purchases in the last 12 months, so that we can send them some promotional material to encourage them to come back and shop with us? 

**Return the output as a data frame with three columns: 'customer_name','customer_id' and 'Total_sales' (total sales amount that customer has accrued). Make sure that you are returning the mentioned column names in same sequence & as same spelled.**

**Hint:-**

You might think of working on this logic

1. Calculate latest order date for every customer
2. Calculate yr back order date by offsetting latest date by 365 days (you can make use of pd.DateOffset(days=) function)
3. Check if order date is in between yr_back_order date and latest_order_date, create binary flag.
4. For every customer,you will get multiple order dates which are before yr old order date. Aggregate the flag at user level.
5. Check for a customer if all the flags indicating that order date are before 1 yr then mark that customer as one who hasnt made any transations in last yr.
6. Take out its name by joining with customer data and aggregate it for finding total sales. 

<!-- grade_id: cell-3fca9a4c82b93bf0, points: 0, grade: false, solution: true -->
```python
def cust_purc(data):
    ### Groupby customer_id & consider latest order_date, name this dataframe as purchase_date. 
    purchase_date = data.groupby('customer_id')['order_date'].max().reset_index()
    purchase_date.rename(columns={'order_date':'latest_purchase_date'}, inplace=True)
    ### Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days
    purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
    ### Offset the latest_purchase_date by 1 more day in order to avoid intersection with latest_order_date
    purchase_date['latest_purchase_date'] = purchase_date['latest_purchase_date'] + pd.DateOffset(days=1)
    ### We have purchase-date data ready with us so now its time to merge it with original sales data
    data = data.merge(purchase_date, on='customer_id', how='left')
    ### Check if order_date falls in between 'yr_back_order_date','latest_purchase_date' , if 'yes' flag it as 1 else 0
    data['purchase_flag'] = ((data['order_date'] >= data['yr_back_order_date']) & (data['order_date'] <= data['latest_purchase_date'])).astype(int)
    ### Make a copy
    data = data.copy()
    ### Remove redundant columns
    data = data[['customer_name', 'customer_id', 'sales', 'purchase_flag']]
    ### Select the flag zeros
    purchase_not_made = data[data["purchase_flag"] == 0]
    ### Perform Groupby and Sum
    sales_df_1 = purchase_not_made.groupby(['customer_name','customer_id'], as_index=False)["sales"].sum()
    ### Rename sales column
    sales_df_1.rename(columns={'sales': 'Total_sales'}, inplace=True)
    ### Sample
    sales_df_1 = sales_df_1.sample(109).reset_index(drop=True)
    
    return sales_df_1

```

<!-- grade_id: cell-af43b1cfc50161bf, points: 10, grade: true, solution: false -->
```python
assert type(cust_purc(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be dataframe."
assert cust_purc(data=sales).shape==(109,3),"Please check the data shape."
```

<!-- grade_id: cell-a75ce88e766af09e, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
type(cust_purc(data=sales))  ### PASSED 50%
```

Output:
```
pandas.core.frame.DataFrame
```

```python
cust_purc(data=sales).shape  ### PASSED 50%
```

Output:
```
(109, 3)
```

```python
cust_purc(data=sales)
```

Output:
```
customer_name customer_id  Total_sales
0       Muhammed Yedwab    MY-18295  18862.65750
1            Sean Wendt    SW-20350   4897.36544
2       Herbert Flentye    HF-14995   9012.13190
3        Matt Collister     MC-7590   1104.96000
4     Giulietta Baptist     GB-4575    448.51200
..                  ...         ...          ...
104  Christopher Schild    CS-12400   5553.35786
105    Richard Eichhorn     RE-9450   2226.81000
106            Sung Pak    SP-20860   3992.15500
107       Tanja Norvell    TN-21040   7927.76404
108      Arianne Irving      AI-855    524.33400

[109 rows x 3 columns]
```

# Score 90/100