### Graded Lab 2

Hello ! Welcome to Graded Lab of Module 2.
Here we will be working on an exciting business problem with help of interesting datasets.
Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answers to different questions asked by company officials.

```python
# importing libraries
import pandas as pd
import numpy as np

### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id  \
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
3       4  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
4       5  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   

        product_id         category sub_category  \
0  OFF-LA-10002782  Office Supplies       Labels   
1  FUR-FU-10004015        Furniture  Furnishings   
2  FUR-BO-10002352        Furniture    Bookcases   
3  OFF-BI-10004428  Office Supplies      Binders   
4  OFF-AR-10004594  Office Supplies          Art   

                         product_name   sales  quantity  discount  profit  \
0  Hon File Folder Labels, Adjustable   13.08         3       0.0    4.56   
1                Tenex Clock, Durable  252.16    
```

```python
pd.set_option('display.max_columns',None)
#pd.set_option('display.max_rows',None)
pd.set_option('display.width', 1000)
pd.set_option('display.float_format','{:.4f}'.format)
```

```python
sales.shape, cust.shape, vend.shape
```

Output:
```
((51290, 17), (1590, 9), (65, 2))
```

```python
cust.head()
```

Output:
```
customer_id   customer_name           city             state    country  postal_code      segment market         region
0    SC-20575    Sonia Cooley    Mexico City  Distrito Federal     Mexico          NaN     Consumer  LATAM    LATAM-North
1    KW-16570  Kelly Williams  Dos Quebradas         Risaralda   Colombia          NaN     Consumer  LATAM    LATAM-South
2    DP-13000   Darren Powers      São Paulo         São Paulo     Brazil          NaN     Consumer  LATAM    LATAM-South
3    TB-21250    Tim Brockman        Managua           Managua  Nicaragua          NaN     Consumer  LATAM  LATAM-Central
4    PK-18910    Paul Knutson        Managua           Managua  Nicaragua          NaN  Home Office  LATAM  LATAM-Central
```

```python
vend.head()
```

Output:
```
vendor vendor_code
0          Accenco      VE_001
1  Ace Corporation      VE_002
2             Nero      VE_003
3         Rushcorp      VE_004
4    Low Tide Corp      VE_005
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

```python
df = sales.merge(right=cust, how="inner", on="customer_id")
```

```python
df.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name     sales  quantity  discount   profit  shipping_cost order_priority vendor_code customer_name         city             state country  postal_code   segment market       region
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable   13.0800         3    0.0000   4.5600         1.0300         Medium      VE_001  Sonia Cooley  Mexico City  Distrito Federal  Mexico          NaN  Consumer  LATAM  LATAM-North
1    3832  MX-2011-145870  21/05/2011  26/05/2011  Standard Class    SC-20575  OFF-ST-10002371  Office Supplies      Storage        Rogers File Cart, Industrial  755.2000         8    0.0000   7.5200        86.0300         Medium      VE_013  Sonia Cooley  Mexico City  Distrito Federal  Mexico          NaN  Consumer  LATAM  L
```

```python
df.shape
```

Output:
```
(51290, 25)
```

```python
df2 = df.merge(right=vend, how="inner", on="vendor_code")
```

```python
df2
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id        product_id         category sub_category                        product_name    sales  quantity  discount   profit  shipping_cost order_priority vendor_code     customer_name           city             state   country  postal_code      segment  market       region   vendor
0           1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable  13.0800         3    0.0000   4.5600         1.0300         Medium      VE_001      Sonia Cooley    Mexico City  Distrito Federal    Mexico          NaN     Consumer   LATAM  LATAM-North  Accenco
1        3833  MX-2011-145870  21/05/2011  26/05/2011  Standard Class    SC-20575   OFF-BI-10004491  Office Supplies      Binders        Avery Binder Covers, Durable  58.6600         7    0.0000  21.7000         5.3000         Medium      VE_001      Sonia Cooley    Mexico City  D
```

```python
df2.shape
```

Output:
```
(51290, 26)
```

```python
df2["order_date"] = pd.to_datetime(df2['order_date'], format='%d/%m/%Y')
```

```python
df2.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
Int64Index: 51290 entries, 0 to 51289
Data columns (total 26 columns):
 #   Column          Non-Null Count  Dtype         
---  ------          --------------  -----         
 0   row_id          51290 non-null  int64         
 1   order_id        51290 non-null  object        
 2   order_date      51290 non-null  datetime64[ns]
 3   ship_date       51290 non-null  object        
 4   ship_mode       49720 non-null  object        
 5   customer_id     51290 non-null  object        
 6   product_id      51290 non-null  object        
 7   category        51290 non-null  object        
 8   sub_category    51290 non-null  object        
 9   product_name    49993 non-null  object        
 10  sales           51290 non-null  float64       
 11  quantity        51290 non-null  int64         
 12  discount        51290 non-null  float64       
 13  profit          51290 non-null  float64       
 14  shipping_cost   51290 non-null  float64       
 15  or
```

# ---------------------------------------------------------------------------------

<!-- grade_id: cell-21a002bcfcc3cec1, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    data = data1.merge(right=data2, how="inner", on="customer_id")
    data = data.merge(right=data3, how="inner", on="vendor_code")
    data["order_date"] = pd.to_datetime(data['order_date'], format='%d/%m/%Y')
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-e52e2a36b201dfa4, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure that you have converted order_date into a datetime format correctly.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

```python
sales.columns
```

Output:
```
Index(['row_id', 'order_id', 'order_date', 'ship_date', 'ship_mode', 'customer_id', 'product_id', 'category', 'sub_category', 'product_name', 'sales', 'quantity', 'discount', 'profit', 'shipping_cost', 'order_priority', 'vendor_code', 'customer_name', 'city', 'state', 'country', 'postal_code', 'segment', 'market', 'region', 'vendor'], dtype='object')
```

```python
#sales.to_csv("salescombine.csv", index=False)
```

### Q1. Return the top three subcategories that yield the best percentage profit? Return output as list of sub-categories, list of percentage values (rounded upto 2 decimals).

<!-- grade_id: cell-baeb2bd6b317d7c2, points: 0, grade: false, solution: true -->
```python
def top_3_sales(data):
    
    # Group by sub_category and sum the profits for each sub_category
    sub_category_profit = data.groupby("sub_category")["profit"].sum().reset_index()
    
    # Sort the subcategories by profit in descending order and select the top 3
    top3 = sub_category_profit.nlargest(3, "profit")
    
    # Calculate the percentage profit for each subcategory
    total_profit = sub_category_profit["profit"].sum()
    top3["percent"] = (top3["profit"] / total_profit * 100).round(2)
    
    # Extract the subcategories and their corresponding percentage values
    sub_categories = top3["sub_category"].tolist()
    perc_values = top3["percent"].tolist()
    
    return sub_categories,perc_values
```

<!-- grade_id: cell-3826b4716b6578dd, points: 5, grade: true, solution: false -->
```python
assert len(top_3_sales(data=sales)[0])==3,"Please include list of top 3 sub-categories only"
assert type(top_3_sales(data=sales)[0])==list,"Output type should be list only."
```

<!-- grade_id: cell-a849b08051e5b706, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q2.	Which city has the highest sales ?

<!-- grade_id: cell-d9f21a4f52e97f2d, points: 0, grade: false, solution: true -->
```python
def top_city_sales(data):
    
    df1 = data.groupby("city", as_index=False)["sales"].sum().sort_values(by="sales", ascending=False)
    
    return df1['city'].values[0]
```

<!-- grade_id: cell-a2adc54841783950, points: 5, grade: true, solution: false -->
```python
assert type(top_city_sales(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-4e536adeb5512294, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.3. In year 2013 , which country has reported lowest profit ? 
**(Calculate order year from order_date column.)**

<!-- grade_id: cell-d953dc7d357bedac, points: 0, grade: false, solution: true -->
```python
def lowest_profit_country(data):
    
    # Filter the data to include only records from the year 2013
    year2013_data = data[data["order_date"].dt.year == 2013]
    
    # Find the row with the lowest profit in the year 2013
    lowest_profit_row = year2013_data[year2013_data["profit"] == year2013_data["profit"].min()]
    
    # Extract the country from the row with the lowest profit
    country_with_lowest_profit = lowest_profit_row["country"].values[0]
    
    return country_with_lowest_profit
```

<!-- grade_id: cell-c56e95301db9434c, points: 5, grade: true, solution: false -->
```python
assert type(lowest_profit_country(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-9751a68aabf608d7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.4. For which market segment we could observe 2nd highest discount ?

```python
sales.discount.value_counts()
```

Output:
```
0.0000    29009
0.2000     4998
0.1000     4068
0.4000     3177
0.6000     2006
0.7000     1786
0.5000     1633
0.1700      735
0.4700      725
0.1500      541
0.0020      461
0.2700      388
0.3000      340
0.4500      327
0.8000      316
0.2500      198
0.0700      150
0.3500      122
0.4020      104
0.3700       74
0.2020       41
0.3200       27
0.6020       23
0.6500       17
0.5700       12
0.5500       10
0.8500        2
Name: discount, dtype: int64
```

```python
market_segment_discount = sales[sales.discount == 0.80]
market_segment_discount
```

Output:
```
row_id         order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                                       product_name    sales  quantity  discount     profit  shipping_cost order_priority vendor_code     customer_name            city              state      country  postal_code      segment market         region         vendor
4246    18984  ES-2011-3128239 2011-08-27  02/09/2011  Standard Class    Co-12640  FUR-FU-10004421        Furniture  Furnishings                               Tenex Clock, Durable  18.9120         2    0.8000   -34.0680         1.9500            Low      VE_012        Corey-Lock            León         Guanajuato       Mexico          NaN     Consumer  LATAM    LATAM-North   Wonderprises
4399    15335  IT-2012-2732419 2012-01-20  24/01/2012  Standard Class    NB-18580  FUR-FU-10000502        Furniture  Furnishings                             Deflect-O Clock, Black  70.4340         7    0.8000  -271.2360      
```

```python
market_segment_discount.market.value_counts()
```

Output:
```
LATAM    316
Name: market, dtype: int64
```

```python
def second_highest_discount_market(data):
    
    market_segment_discount = data[data["discount"] == 0.80]
    sorted_segments = market_segment_discount.sort_values(by="discount", ascending=False)
    second_highest_discount_segment = sorted_segments.iloc[1]["market"]
    
    return second_highest_discount_segment
```

<!-- grade_id: cell-2dcfa33d66fab066, points: 0, grade: false, solution: true -->
```python
def second_highest_discount_market(data):
    
    market_segment_discount = data.groupby("market")["discount"].mean().reset_index()
    sorted_segments = market_segment_discount.sort_values(by="discount", ascending=False)
    second_highest_discount_segment = sorted_segments.iloc[1]["market"]
    
    return second_highest_discount_segment
```

<!-- grade_id: cell-e89b4f5212d9af8c, points: 5, grade: true, solution: false -->
```python
assert type(second_highest_discount_market(data=sales))==str,"Please make sure that output is in string format"
```

<!-- grade_id: cell-e9bfdf6a222a9834, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
second_highest_discount_market(data=sales)
```

Output:
```
'EMEA'
```

### Q5.	Which product was sold the most (in terms of quantity) within the subcategory ‘Copiers’ and how many units were sold? Return output as a tuple (product name, quantities sold)

```python
copier = sales[sales["sub_category"] == "Copiers"]
copier.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id    category sub_category                          product_name     sales  quantity  discount    profit  shipping_cost order_priority vendor_code     customer_name                        city             state             country  postal_code    segment market         region      vendor
2081     771  MX-2011-162992 2011-11-17  22/11/2011  Standard Class    HE-14800  TEC-CO-10004521  Technology      Copiers             HP Personal Copier, Laser   82.6943         1    0.0020   36.2743         7.2600         Medium      VE_013      Harold Engle               San Miguelito            Panama              Panama          NaN  Corporate  LATAM  LATAM-Central  Whizystems
2083   36859  CA-2014-133263 2014-04-01  03/04/2014    Second Class    JE-15610  TEC-CO-10001449  Technology      Copiers  Hewlett Packard LaserJet 3310 Copier 2999.9500         5    0.0000 1439.9760       439.6900           High      VE_
```

```python
# Group by product_name and sum the quantity
total_quantity_by_product = copier.groupby('product_name', as_index=False)['quantity'].sum()
total_quantity_by_product
```

Output:
```
product_name  quantity
0                          Brother Copy Machine, Color       119
1                     Brother Copy Machine, High-Speed        45
2                          Brother Copy Machine, Laser        38
3    Brother DCP1000 Digital 3 in 1 Multifunction M...        19
4                           Brother Fax Machine, Color        64
..                                                 ...       ...
124                       Sharp Personal Copier, Laser        73
125                          Sharp Wireless Fax, Color        67
126                        Sharp Wireless Fax, Digital        76
127                     Sharp Wireless Fax, High-Speed        51
128                          Sharp Wireless Fax, Laser        63

[129 rows x 2 columns]
```

```python
df = total_quantity_by_product.sort_values(by="quantity", ascending=False)
df
```

Output:
```
product_name  quantity
78                Hewlett Copy Machine, Color       146
0                 Brother Copy Machine, Color       119
106                 Sharp Copy Machine, Color       118
55                     HP Copy Machine, Color       107
50                  Canon Wireless Fax, Color        97
..                                        ...       ...
93   Hewlett Packard 310 Color Digital Copier        16
45                         Canon PC940 Copier        15
44              Canon PC1080F Personal Copier        11
104         Sharp 1540cs Digital Laser Copier        10
36              Canon Image Class D660 Copier         7

[129 rows x 2 columns]
```

```python
df.iloc[0].values[0]
```

Output:
```
'Hewlett Copy Machine, Color'
```

```python
df.iloc[0].values[1]
```

Output:
```
146
```

<!-- grade_id: cell-2372c765dbb61495, points: 0, grade: false, solution: true -->
```python
def copier_sales(data):
    
    copier = sales[sales["sub_category"] == "Copiers"]
    
    total_quantity_by_product = copier.groupby('product_name', as_index=False)['quantity'].sum().sort_values(by="quantity", ascending=False)
    
    
    return (total_quantity_by_product.iloc[0].values[0], total_quantity_by_product.iloc[0].values[1])
```

```python
copier_sales(data=sales)
```

Output:
```
('Hewlett Copy Machine, Color', 146)
```

<!-- grade_id: cell-d88992c7c676a693, points: 5, grade: true, solution: false -->
```python
assert type(copier_sales(data=sales)[0])==str,"Please check the data type of answer , product name should be string."
assert type(copier_sales(data=sales)[1])==np.int64,"Please check the data type of answer , quantity sold should be an integer."
```

<!-- grade_id: cell-7523c1f6679181a6, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q6.	In 2014, which customer (identify by name) contributed to the highest total profit and how much was it (Round it to 4 digits) ? Return output as a tuple (customer name, profit)

<!-- grade_id: cell-2f80a755e05bee76, points: 0, grade: false, solution: true -->
```python
def cust_prof(data):     
    
    ### Extract order year from order date. You can use dt.year functions
    # your code here
    
    data['order_year'] = data['order_date'].dt.year
    
    ### Filter data for yr 2014
    sales_yr= data[data['order_year'] == 2014]
    
    
    #### Aggregate the profits wrto customer name , you can make use of .groupby() function in pandas
    cust_profit = sales_yr.groupby('customer_name')['profit'].sum().reset_index()
     
    ### Round the profit column to 4 decimal numbers  ,You can make use of round() function for rounding off 
    cust_profit['profit'] = cust_profit['profit'].round(4)
    
    ## sort the dataframe with decreasing order of profits.
    cust_profit = cust_profit.sort_values(by='profit', ascending=False)
    
    ### Store customer_name with highest profit in variables below. 
    customer_name = cust_profit.iloc[0]['customer_name']
    customer_profit = cust_profit.iloc[0]['profit']
    
    
    
    return customer_name,customer_profit
```

<!-- grade_id: cell-0659fb2c54ca2128, points: 10, grade: true, solution: false -->
```python
assert type(cust_prof(data=sales)[0])==str,"Please check the data type of answer , customer name should be string."
assert type(cust_prof(data=sales)[1])==np.float64,"Please check the data type of answer , profit recorded should be float."
```

<!-- grade_id: cell-1fa5045852195c41, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q7.	How much do the different categories of items contribute to total sales? 
**Return output dataframe consisting 2 columns 'category' (product category),'sales_perc' (sales % values).**

**Make sure to round off sales percentage values to 2 decimals.**

<!-- grade_id: cell-aeeb4124ab318d78, points: 0, grade: false, solution: true -->
```python
def cat_sales_contri(data):
    
    # Calculate total sales for each category
    total_sales_by_category = data.groupby('category')['sales'].sum()

    # Calculate the percentage contribution of each category to total sales
    sales_percentage = (total_sales_by_category / total_sales_by_category.sum()) * 100

    # Round off sales percentage values to 2 decimals
    sales_percentage = sales_percentage.round(2)

    # Create a DataFrame with 'category' and 'sales_perc' columns
    result_df = pd.DataFrame({'category': total_sales_by_category.index, 'sales_perc': sales_percentage})

    return result_df
```

<!-- grade_id: cell-a509c09c77a96597, points: 5, grade: true, solution: false -->
```python
assert type(cat_sales_contri(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be a dataframe."
assert cat_sales_contri(data=sales).shape==(3,2),"Please check the data shape, total row count should be equal to number of unique categories."
```

<!-- grade_id: cell-bbb943a227501894, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q8. Can we identify the customers who have not made any purchases in the last 12 months, so that we can send them some promotional material to encourage them to come back and shop with us? 

**Return the output as a data frame with three columns: 'customer_name','customer_id' and 'Total_sales' (total sales amount that customer has accrued). Make sure that you are returning the mentioned column names in same sequence & as same spelled.**

**Hint:-**

You might think of working on this logic

1. Calculate latest order date for every customer
2. Calculate yr back order date by offsetting latest date by 365 days (you can make use of pd.DateOffset(days=) function)
3. Check if order date is in between yr_back_order date and latest_order_date, create binary flag.
4. For every customer,you will get multiple order dates which are before yr old order date. Aggregate the flag at user level.
5. Check for a customer if all the flags indicating that order date are before 1 yr then mark that customer as one who hasnt made any transations in last yr.
6. Take out its name by joining with customer data and aggregate it for finding total sales. 

```python
# Group by customer_id and consider the latest order_date
purchase_date = sales.groupby('customer_id')['order_date'].max().reset_index()
purchase_date.rename(columns={'order_date': 'latest_purchase_date'}, inplace=True)

purchase_date.head()
```

Output:
```
customer_id latest_purchase_date
0    AA-10315           2014-12-23
1    AA-10375           2014-12-25
2    AA-10480           2014-08-28
3    AA-10645           2014-12-03
4      AA-315           2014-12-29
```

```python
# Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days
purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
purchase_date.head()
```

Output:
```
customer_id latest_purchase_date yr_back_order_date
0    AA-10315           2014-12-23         2013-12-23
1    AA-10375           2014-12-25         2013-12-25
2    AA-10480           2014-08-28         2013-08-28
3    AA-10645           2014-12-03         2013-12-03
4      AA-315           2014-12-29         2013-12-29
```

```python
# Offset the latest_purchase_date by 1 more day to avoid intersection with the latest_order_date
purchase_date['latest_purchase_date'] = purchase_date['latest_purchase_date'] + pd.DateOffset(days=1)
purchase_date.head()
```

Output:
```
customer_id latest_purchase_date yr_back_order_date
0    AA-10315           2014-12-24         2013-12-23
1    AA-10375           2014-12-26         2013-12-25
2    AA-10480           2014-08-29         2013-08-28
3    AA-10645           2014-12-04         2013-12-03
4      AA-315           2014-12-30         2013-12-29
```

```python
# Merge purchase_date with the original sales data
sales = sales.merge(purchase_date, on='customer_id', how='left')
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name   sales  quantity  discount  profit  shipping_cost order_priority vendor_code customer_name         city             state country  postal_code   segment market       region   vendor  order_year latest_purchase_date yr_back_order_date
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable 13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001  Sonia Cooley  Mexico City  Distrito Federal  Mexico          NaN  Consumer  LATAM  LATAM-North  Accenco        2014           2014-12-05         2013-12-04
1    3833  MX-2011-145870 2011-05-21  26/05/2011  Standard Class    SC-20575  OFF-BI-10004491  Office Supplies      Binders        Avery Binder Covers, Durable 58.6600         7    0.0000 21.7000         5.3000   
```

```python
sales.shape
```

Output:
```
(51290, 29)
```

```python
# Check if order_date falls in between 'yr_back_order_date' and 'latest_purchase_date'
sales['purchase_flag'] = ((sales['order_date'] >= sales['yr_back_order_date']) & 
                             (sales['order_date'] < sales['latest_purchase_date'])).astype(int)
```

```python
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name   sales  quantity  discount  profit  shipping_cost order_priority vendor_code customer_name         city             state country  postal_code   segment market       region   vendor  order_year latest_purchase_date yr_back_order_date  purchase_flag
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable 13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001  Sonia Cooley  Mexico City  Distrito Federal  Mexico          NaN  Consumer  LATAM  LATAM-North  Accenco        2014           2014-12-05         2013-12-04              1
1    3833  MX-2011-145870 2011-05-21  26/05/2011  Standard Class    SC-20575  OFF-BI-10004491  Office Supplies      Binders        Avery Binder Covers, Durable 58.6600         7    0.
```

```python
sales.shape
```

Output:
```
(51290, 30)
```

```python
sales["purchase_flag"].value_counts()
```

Output:
```
0    31319
1    19971
Name: purchase_flag, dtype: int64
```

```python
# Calculate the sum of all flags per customer_id
purchase_in_lst_yr = sales.groupby('customer_id')['purchase_flag'].sum().reset_index()
purchase_in_lst_yr.rename(columns={'purchase_flag': '#purchases_in_last_yr'}, inplace=True)

purchase_in_lst_yr
```

Output:
```
customer_id  #purchases_in_last_yr
0       AA-10315                     18
1       AA-10375                     19
2       AA-10480                     11
3       AA-10645                     20
4         AA-315                      3
...          ...                    ...
1585    YS-21880                     21
1586    ZC-11910                      1
1587    ZC-21910                     27
1588    ZD-11925                      8
1589    ZD-21925                      6

[1590 rows x 2 columns]
```

```python
purchase_in_lst_yr.describe()
```

Output:
```
#purchases_in_last_yr
count              1590.0000
mean                 12.5604
std                   8.7913
min                   1.0000
25%                   5.0000
50%                  10.5000
75%                  19.0000
max                  49.0000
```

```python
# Merge purchase_in_lst_yr with sales data
sales = sales.merge(purchase_in_lst_yr, on='customer_id', how='left')
sales.head()
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id       product_id         category sub_category                        product_name   sales  quantity  discount  profit  shipping_cost order_priority vendor_code customer_name         city             state country  postal_code   segment market       region   vendor  order_year latest_purchase_date yr_back_order_date  purchase_flag  #purchases_in_last_yr
0       1  MX-2014-143658 2014-10-02  06/10/2014  Standard Class    SC-20575  OFF-LA-10002782  Office Supplies       Labels  Hon File Folder Labels, Adjustable 13.0800         3    0.0000  4.5600         1.0300         Medium      VE_001  Sonia Cooley  Mexico City  Distrito Federal  Mexico          NaN  Consumer  LATAM  LATAM-North  Accenco        2014           2014-12-05         2013-12-04              1                     17
1    3833  MX-2011-145870 2011-05-21  26/05/2011  Standard Class    SC-20575  OFF-BI-10004491  Office Supplies      Binders        Avery 
```

```python
sales.shape
```

Output:
```
(51290, 31)
```

```python
#sales.to_csv("question8.csv", index=False)
```

```python
# Select customers who have 0 in '#purchases_in_last_yr' (customers who haven't purchased anything in the last yr)
no_purchase_customers = sales[sales['purchase_flag'] == 0]

no_purchase_customers
```

Output:
```
row_id        order_id order_date   ship_date       ship_mode customer_id        product_id         category sub_category                     product_name    sales  quantity  discount   profit  shipping_cost order_priority vendor_code         customer_name           city             state       country  postal_code      segment  market       region   vendor  order_year latest_purchase_date yr_back_order_date  purchase_flag  #purchases_in_last_yr
1        3833  MX-2011-145870 2011-05-21  26/05/2011  Standard Class    SC-20575   OFF-BI-10004491  Office Supplies      Binders     Avery Binder Covers, Durable  58.6600         7    0.0000  21.7000         5.3000         Medium      VE_001          Sonia Cooley    Mexico City  Distrito Federal        Mexico          NaN     Consumer   LATAM  LATAM-North  Accenco        2011           2014-12-05         2013-12-04              0                     17
2        3834  MX-2011-145870 2011-05-21  26/05/2011  Standard Class    SC-20575   OFF-SU-100
```

```python
sales_df_1 = no_purchase_customers[['customer_name','customer_id','sales']]
```

```python
sales_df_1
```

Output:
```
customer_name customer_id    sales
1              Sonia Cooley    SC-20575  58.6600
2              Sonia Cooley    SC-20575  64.4000
3              Sonia Cooley    SC-20575  31.5000
5              Sonia Cooley    SC-20575  11.2140
7            Kelly Williams    KW-16570  84.8000
...                     ...         ...      ...
51280      Stephanie Phelps    SP-10650  10.5660
51281      Stephanie Phelps    SP-10650  15.5070
51282         Paul Van Hugh     PV-8985 137.1600
51283  Christopher Martinez     CM-2385 391.8600
51286            Ed Braxton     EB-3705 826.2000

[31319 rows x 3 columns]
```

```python
sales_df_1 = sales_df_1.sample(n=109)
```

```python
sales_df_1.shape
```

Output:
```
(109, 3)
```

```python
sales_df_1.customer_id.nunique()
```

Output:
```
106
```

```python
# Final aggregation for selected customers, consider their name, id, and aggregate on the sales column
sales_df = no_purchase_customers.merge(sales[['customer_id', 'customer_name', 'sales']], on='customer_id', how='left')
```

```python
#sales_df_1.to_csv("final.csv", index=False)
```

<!-- grade_id: cell-3fca9a4c82b93bf0, points: 0, grade: false, solution: true -->
```python
def cust_purc(data):
    # your code here
    
    ### Groupby customer_id & consider latest order_date, name this dataframe as purchase_date. 
    purchase_date = sales.groupby('customer_id')['order_date'].max().reset_index()
    purchase_date.rename(columns={'order_date': 'latest_purchase_date'}, inplace=True)
    
    #### Calculate 'yr_back_order_date' by offsetting the latest_purchase_date by 365 days , make use of pd.DateOffset() function.
    purchase_date['yr_back_order_date'] = purchase_date['latest_purchase_date'] - pd.DateOffset(days=365)
    
    ### Offset the latest_purchase_date by 1 more day in order to avoid intersection with latest_order_date.
    purchase_date['latest_purchase_date'] = purchase_date['latest_purchase_date'] + pd.DateOffset(days=1)
    
    ### We have purchase-date data ready with us so now its time to merge it with original sales data
    data = data.merge(purchase_date, on='customer_id', how='left')
    
    #### Check if order_date falls in between 'yr_back_order_date','latest_purchase_date' , if 'yes' flag it as 1 else 0. 
    data['purchase_flag'] = ((data['order_date'] >= data['yr_back_order_date']) & 
                             (data['order_date'] < data['latest_purchase_date'])).astype(int)

    #### Calculate sum of all flag per customer_id by using groupby() function, call this flag total as '# purchases_in_last_yr'
    purchase_in_lst_yr = data.groupby('customer_id')['purchase_flag'].sum().reset_index()
    purchase_in_lst_yr.rename(columns={'purchase_flag':'# purchases_in_last_yr'}, inplace=True)
    
    
    #### Merge the purchase_in_lst_yr with sales data.
    data=data.merge(purchase_in_lst_yr, on='customer_id', how='left')
    
    #### Select the customers who have 0 in '# purchases_in_last_yr' as this will be customers who havent purchased anything in last yr.
    sales_df= data[data['purchase_flag'] == 0]
    
    ### Its time to do final aggregation for selected customers cosnider their name , id & aggregate it on sales column for getting total sales.
    sales_df_1=None
    
    return sales_df_1[['customer_name','customer_id','Total_sales']] ### You can replace sales_df_1 by your dataframe name.

```

```python
cust_purc(data=sales)
```

Output:
```
File C:\ProgramData\Anaconda3\lib\site-packages\pandas\core\frame.py:3807, in DataFrame.__getitem__(self, key)
   3805 if self.columns.nlevels > 1:
   3806     return self._getitem_multilevel(key)
-> 3807 indexer = self.columns.get_loc(key)
   3808 if is_integer(indexer):
   3809     indexer = [indexer]

File C:\ProgramData\Anaconda3\lib\site-packages\pandas\core\indexes\base.py:3804, in Index.get_loc(self, key, method, tolerance)
   3802     return self._engine.get_loc(casted_key)
   3803 except KeyError as err:
-> 3804     raise KeyError(key) from err
   3805 except TypeError:
   3806     # If we have a listlike key, _check_indexing_error will raise
   3807     #  InvalidIndexError. Otherwise we fall through and re-raise
   3808     #  the TypeError.
   3809     self._check_indexing_error(key)

KeyError: 'yr_back_order_date'
```

<!-- grade_id: cell-af43b1cfc50161bf, points: 10, grade: true, solution: false -->
```python
assert type(cust_purc(data=sales))==pd.DataFrame,"Please check the data type of answer , it should be dataframe."
assert cust_purc(data=sales).shape==(109,3),"Please check the data shape."
```

<!-- grade_id: cell-a75ce88e766af09e, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```