### Graded Lab 3

Hello ! Welcome to Graded Lab of Module 3.

Working on the same analysis task.

* Here we will work on correlations between variables , testing few of the hypothesis business wants us to test.

In case if you are not able to recollect the problem description and data description then mentioning it below.

Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answers to different questions asked by company officials.

# Instructor update: Dec 14, 2023

```python
# importing libraries
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from scipy.stats import norm
from itertools import combinations
from scipy import stats
from statsmodels.stats import weightstats as ztests


### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-d6f0b3d1d3b16d3f, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    # your code here
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-6e7272d20447184e, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure that you have converted order_date into a datetime format correctly.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

### Q1. What is skewness,skewness type ,kurtosis, kurtosis type for profit column ? 
#### Return output in tuple format.Eg :- (-1, "Left-skewed", 2, "Platykurtic"). Make sure to round off answer to 2 decimals.


Note:-
1. Skewness type takes 3 values :- Right-skewed (Skewness >0) ,Left-skewed (Skewness <0), Symmetric (Skewness = 0)
2. Kurtosis type takes 3 values :- Platykurtic (kurtosis < 3) ,Leptokurtic (kurtosis > 3), Mesokurtic (kurtosis = 3)

<!-- grade_id: cell-15b38d6aecccc0d5, points: 0, grade: false, solution: true -->
```python
def skew_kurt(data,column):    
    # your code here
```

<!-- grade_id: cell-484293228e575f1b, points: 5, grade: true, solution: false -->
```python
assert type(skew_kurt(data=sales,column='profit'))==tuple, "Please provide output in a tuple format."
assert len(skew_kurt(data=sales,column='profit'))==4,"Please include skewness , skewness_type , kurtosis & kurtosis_type."
assert skew_kurt(data=sales,column='profit')[1] in ["Right-skewed","Left-skewed","Symmetric"],"Please include from the listed types only , also make sure that spelling is correct."
assert skew_kurt(data=sales,column='profit')[3] in ["Platykurtic","Leptokurtic","Mesokurtic"],"Please include from the listed types only , also make sure that spelling is correct."
```

<!-- grade_id: cell-69fee297270967e0, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q2.	What is correlation between different numerical variables ? ('sales', 'quantity', 'discount', 'profit', 'shipping_cost')
### Return a dataframe with 3 columns : Variable 1 , Variable 2 , Correlation. 
### Round off correlation values upto 4 decimals.
### Sort dataframe with respect to highest absolute correlation & make sure that one variable pair occurs only once. For eg : (Col A , Col B) has occured then (Col B , Col A) cant occur.

<!-- grade_id: cell-891376cce26eb132, points: 0, grade: false, solution: true -->
```python
def corr_sales(data):
    # your code here
    return pair_df[['variable_1','variable_2','corr']]
```

<!-- grade_id: cell-d39bc7b2a3f0e000, points: 7, grade: true, solution: false -->
```python
assert type(corr_sales(data=sales))==pd.DataFrame, "Please provide output in dataframe format."
assert all(corr_sales(data=sales).columns==['variable_1', 'variable_2', 'corr']), "Make sure your output columns are named as 'variable_1', 'variable_2', 'corr'."
assert corr_sales(data=sales).shape[1]==3, "Make sure you are returning only 3 columns."
assert corr_sales(data=sales).shape[0]==10, "Make sure you are returning all possible combinations."
assert list(corr_sales(data=sales)['corr'])[0]== 0.7681, "Please provide correct value , make sure that answer is rounded upto 4 digits."
assert list(corr_sales(data=sales)['corr'])[-1]== -0.0199, "Please provide correct value , make sure that answer is rounded upto 4 digits."
```

<!-- grade_id: cell-e680b68e89f9759a, points: 3, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q3.	 Plot a box plot for sales in 'Furniture' category and tell on which side of box-plot you were able to see outliers.

### Store the output in 'outliers_whiskers' variable. Output should be from one of the following choices : 'both_sides' , 'upper_whisker' ,'lower_whisker'. 

<!-- grade_id: cell-c3127d529f222152, points: 0, grade: false, solution: true -->
```python
def outlier_sales(data): 
    # your code here
    return outliers_whiskers # outlier_whsikers is a variable in where you will store answer selected from this 3 options : both_sides , upper_whisker , lower_whisker
```

<!-- grade_id: cell-aae54343919bbde4, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q4.	What is probability that profit lies between 13 & 25.
Return output as float type (check if output dtype is numpy.float64) ,rounded upto 4 decimals.

<!-- grade_id: cell-91216071cd4b7a14, points: 0, grade: false, solution: true -->
```python
def prob(data,lower_bound,upper_bound):
    # your code here
```

<!-- grade_id: cell-b4e6e62544e5e4c0, points: 5, grade: true, solution: false -->
```python
assert type(prob(data=sales,lower_bound=13,upper_bound=25))==np.float64, 'Make sure that you are returning output as float64 (check if output dtype is numpy.float64) type , it should be rounded up to 4 decimals.'
assert (prob(data=sales,lower_bound=13,upper_bound=25)>=0) & (prob(data=sales,lower_bound=13,upper_bound=25)<=1), 'Remember probability always lies between 0 to 1.'
```

<!-- grade_id: cell-10113b1e8f9ae5f7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.

```

###  Q.5) Company's Sales team found an interesting insight from past study. Study said that on a 3-yr scale  avg quantities sold comes out to be 3. Now they want to test if the average quantities sold number has increased or not i.e they wish to understand if average quantities sold are greater than a previously known value i.e 3 or not , Our data is on  3 yrs horizon. So they have asked you to draw a sample of size 100 from sales data, analyze and report findings. Select alpha (level of significance) as 0.05.

**Return output as tuple consisting of  (statistics_value, p-value, decision). Make sure to round up statistics_value & p-value to 4 decimals. Your code will be tested on an input sample (sample_1).**


For eg :- (1.267,0.045,'Reject Null hypothesis')

Where ,stastics:- It is calculated statistic from appropriate test.

p-value :- probability value

Decision :- You have to select from the 2 choices 'Reject Null Hypothesis' , 'Do not Reject Null Hypothesis''

**Hint:**

We are trying to test a hypothesis whether average quantities data for sales data are greater than previously known value 3 or not ?
In order to solve this problem , Think on the lines below.

H0: Average quantities sold = previously_known_population_value (3)

H1: Average quantities sold > previously_known_population_value (3) 

We will use 1 sample z-test with alernative type as 'larger'  for solving this problem.

```python
# Please do not alter/ delete /edit this cell.

from scipy import stats
from statsmodels.stats.weightstats import ztest as ztest
import random
random.seed(9001)

sampleSize=100
group_A= sales['quantity']
sample_1 =  random.sample([x for x in group_A], sampleSize)

## sample for sold quantities is given as sample_1
print(f"length of sold 'quantity' sample is {len(sample_1)}")
```

<!-- grade_id: cell-9284ea7b5f59aada, points: 0, grade: true, solution: false -->
```python
assert np.mean(sample_1)==3.42 , "If the sales dataframe is not being created properly then sample mean will not match ,check the creation of sales datafarme again."
```

<!-- grade_id: cell-64ca11743fc74606, points: 0, grade: false, solution: true -->
```python
def one_sample_ztest(sample1,alpha,popn_value):
    ## sample1(quantity) will be a list of values with length 100.
    # your code here
    return round(z_statistic,4),round(p_value,4),decision  
```

<!-- grade_id: cell-f949706e3f6b68fc, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.6)  Sales team is interested in knowing if furniture & office supply products sell equally or is there any differece between their sells so they ask us to draw a sample of size 22 from furniture & sample of size 25 from Office Supplies from sales dataset and report findings. Select alpha (level of significance) as 0.05.

**Return output as tuple consisting of (statistics_value, p-value, decision). Make sure that stastics,p-value are rounded upto 4 digits.Your code will be tested on 2 input samples (sample_1,sample_2).**

For eg :- (1.267,0.045,'Reject Null hypothesis')

Where ,stastics:- It is calculated statistic from chosen test.

p-value :- probability value

Decision :- 'Reject Null Hypothesis' , ''Do not Reject Null Hypothesis''



```python
# Please do not alter/ delete /edit this cell.

import random
random.seed(9001)


sampleSize=22
group_A= sales[sales['category']=='Furniture']['sales']
sample_1 =  random.sample([x for x in group_A], sampleSize)



sampleSize=25
group_B= sales[sales['category']=='Office Supplies']['sales']
sample_2 =  random.sample([x for x in group_B], sampleSize)


## sample for Furniture sales is given as sample_1
print(f'length of Furniture sales sample is {len(sample_1)}')
## sample for Office Supply sales is given as sample_2
print(f'length of Office Supply sales sample is {len(sample_2)}')
```

<!-- grade_id: cell-16e7cce812f1dd10, points: 0, grade: true, solution: false -->
```python
assert round(np.mean(sample_1),2)==367.13 , "If the sales dataframe is not being created properly then sample mean will not match ,check the creation of sales datafarme again."
assert round(np.mean(sample_2),2)==92.96 , "If the sales dataframe is not being created properly then sample mean will not match ,check the creation of sales datafarme again."
```

<!-- grade_id: cell-ea9ba7ce1cdf4790, points: 0, grade: false, solution: true -->
```python
def hypo_test(sample1,sample2,alpha):
    ## sample1(furniture sales) will be a list of values with length 22.
    ## sample2(office supplies sales) will be a list of values with length 25.
    # your code here
    return round(statistic,4),round(p_value,4),decision
```

<!-- grade_id: cell-6884a77818761b56, points: 13, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 1
```

<!-- grade_id: cell-eeb975b0e1b969db, points: 13, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 2
```

<!-- grade_id: cell-b4a5ed1c86e547f1, points: 4, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 2
```

## Great Job ! ! !
Business team is happy with your work they wish to test couple of more things & require your help.

### Q.7)  Sales team is interested in knowing if profits in EMEA & Africa region are same or not. Test the hypothesis. Select alpha (level of significance) as 0.05.

**Return output as tuple consisting of (statistics_value,p-value,decision). Make sure that stastics,p-value are rounded upto 4 digits.Your code will be tested on 2 input samples (sample_1,sample_2).**


For eg :- (1.267,0.045,'Reject Null hypothesis')

Where ,statistics_value:- It is calculated statistic from chosen test.

p-value :- probability value

Decision :- 'Reject Null Hypothesis' , ''Do not Reject Null Hypothesis''



```python
# Please do not alter/ delete /edit this cell
import random
random.seed(9001)

sampleSize=500
group_A= sales[sales['region']=='EMEA']['profit']
sample_1 =  random.sample([x for x in group_A], sampleSize)

sampleSize=567
group_B= sales[sales['region']=='Africa']['profit']
sample_2 =  random.sample([x for x in group_B], sampleSize)

## sample for EMEA is given as sample_1
print(f'length of EMEA sample is {len(sample_1)}')
## sample for Africa is given as sample_2
print(f'length of Africa sample is {len(sample_2)}')
```

<!-- grade_id: cell-00857b7d8bcf1ec1, points: 0, grade: true, solution: false -->
```python
assert round(np.mean(sample_1),2)==13.45 , "If the sales dataframe is not being created properly then sample mean will not match ,check the creation of sales datafarme again."
assert round(np.mean(sample_2),2)==27.44 , "If the sales dataframe is not being created properly then sample mean will not match ,check the creation of sales datafarme again."
```

<!-- grade_id: cell-9a6efa7bc0d65e78, points: 0, grade: false, solution: true -->
```python
def hypo_test(sample1,sample2,alpha):
    ## sample_1 (EMEA region) will be a list of values with length 500.
    ## sample_2 (Africa region) will be a list of values with length 567.
    # your code here
    return round(statistic,4),round(p_value,4),decision
```

<!-- grade_id: cell-33092e29128eb474, points: 8, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 1
```

<!-- grade_id: cell-7e5ee2d1fd5220a4, points: 9, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 2
```

<!-- grade_id: cell-ced6f05e544babe6, points: 3, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
### Hidden test 3
```