### Graded Lab 4

Hello ! Welcome to Graded Lab of Module 4.

In the last assignment we had worked on correlation , hypothesis which business wanted us to test.
Its high time when we look at the missing values & outliers stuff.
**Here we will work on missing values & outliers.**

In case if you are not able to recollect the problem description and data description then mentioning it below.

Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answer to different questions asked by company officials.

```python
# importing libraries
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from scipy.stats import norm
from itertools import combinations
from sklearn.preprocessing import OrdinalEncoder, LabelEncoder

### Reading sales data
sales = pd.read_csv('sales_data_M4.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')

sales.head()
```

Output:
```
row_id        order_id  order_date   ship_date       ship_mode customer_id  \
0       1  MX-2014-143658  02/10/2014  06/10/2014  Standard Class    SC-20575   
1       2  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
2       3  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
3       4  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   
4       5  MX-2012-155047  15/10/2012  20/10/2012  Standard Class    KW-16570   

        product_id         category sub_category  \
0  OFF-LA-10002782  Office Supplies       Labels   
1  FUR-FU-10004015        Furniture  Furnishings   
2  FUR-BO-10002352        Furniture    Bookcases   
3  OFF-BI-10004428  Office Supplies      Binders   
4  OFF-AR-10004594  Office Supplies          Art   

                         product_name   sales  quantity  discount  profit  \
0  Hon File Folder Labels, Adjustable   13.08         3       0.0    4.56   
1                Tenex Clock, Durable  252.16    
```

```python
sales.shape
```

Output:
```
(51290, 17)
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-760043ec960a09cb, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    data = data1.merge(right=data2, how="inner", on="customer_id")
    data = data.merge(right=data3, how="inner", on="vendor_code")
    data["order_date"] = pd.to_datetime(data['order_date'])
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

Output:
```
C:\Users\Dennis\AppData\Local\Temp\ipykernel_6376\3097288081.py:6: UserWarning: Parsing dates in DD/MM/YYYY format when dayfirst=False (the default) was specified. This may lead to inconsistently parsed dates! Specify a format to ensure consistent parsing.
  data["order_date"] = pd.to_datetime(data['order_date'])
```

<!-- grade_id: cell-24e204e9afc2c313, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure if you have converted order_date into a datetime format correctly or not.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

Note that in this dataset we are purposefully introducing some of the missing values.

### Q.1 Return column names & missing values percentage for the columns which have % misisng value >50. Output should be a dictionary. For eg:- { column_name : %}, make sure to round off percentages to 2 decimals.

<!-- grade_id: cell-eaa4ef1cb3c93902, points: 0, grade: false, solution: true -->
```python
def missing_value_col(data):
    
    # Calculate the percentage of missing values for each column
    missing_percentages = data.isnull().mean() * 100

    # Filter out columns which have more than 50% missing values
    missing_percentages = {col: round(percent, 2) for col, percent in missing_percentages.items() if percent > 50}

    return missing_percentages
```

<!-- grade_id: cell-6d0d4c99764805c4, points: 10, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
missing_value_col(sales)
```

Output:
```
{'postal_code': 100.0}
```

#### **Senior data scientists after consulting with the business have decided that such columns needs to be dropped before going further.**

```python
high_missing_col=list(missing_value_col(data=sales).keys()) #Enter the columns with high % of missing values
### We will drop columns with such high missing values
sales.drop(high_missing_col,axis=1,inplace=True)
```

```python
list(missing_value_col(data=sales).keys())
```

Output:
```
[]
```

```python
sales.columns
```

Output:
```
Index(['row_id', 'order_id', 'order_date', 'ship_date', 'ship_mode',
       'customer_id', 'product_id', 'category', 'sub_category', 'product_name',
       'sales', 'quantity', 'discount', 'profit', 'shipping_cost',
       'order_priority', 'vendor_code', 'customer_name', 'city', 'state',
       'country', 'segment', 'market', 'region', 'vendor'],
      dtype='object')
```

### Q.2 Fill the missing data in 'category' column with appropriate method.  
(Hint:- Remember variable 'category' & 'sub-category' are related & hierachy goes like this category -> sub-category)

Approach:-

1. Create a key-pair dictionary where key will be product category & pair values will be sub-categories.

2. Wherever category value is missing check for corresponding sub-category value. 

3. From sub-category value trace back its original product category value by looking in dictionary, get a product value.

4. Fill the blank product category value by the same.

<!-- grade_id: cell-21490e58b10fdc96, points: 0, grade: false, solution: true -->
```python
### create a dictionary where keys will be categories & sub-categories as values. store in inside Category_dict.

def cat_dict(data,main_col,sub_col):    
    
    # Create a dictionary where keys are categories and values are lists of sub-categories
    Category_dict = data.groupby(main_col)[sub_col].apply(list).to_dict()
    # Remove duplicates in each list
    Category_dict = {k: list(set(v)) for k, v in Category_dict.items()}
    
    return Category_dict 

def category_subcategory_map_dataframe(df, Category_dict, main_col, sub_col):
    def category_subcategory_map(row):
        if pd.isna(row[main_col]):  # If the main column is null
            if not pd.isna(row[sub_col]):  # If the sub column is not null
                for k, v in Category_dict.items():
                    if row[sub_col] in Category_dict[k]:
                        return k  # k is the correct sub-category for the category value v
        else:
            return row[main_col]  # If the main column is not null, return its existing value

    df[main_col] = df.apply(category_subcategory_map, axis=1)
    return df[main_col]
```

```python
Category_dict=cat_dict(data=sales,main_col='category',sub_col='sub_category')
sales['category']=category_subcategory_map_dataframe(df=sales, Category_dict=Category_dict, main_col='category', sub_col='sub_category')
```

<!-- grade_id: cell-b7f7ba970eee54ec, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
cat_dict(sales,main_col="category", sub_col="sub_category")
```

Output:
```
{'Furniture': ['Furnishings', 'Chairs', 'Bookcases', 'Tables'],
 'Office Supplies': ['Storage',
  'Supplies',
  'Appliances',
  'Fasteners',
  'Paper',
  'Art',
  'Envelopes',
  'Labels',
  'Binders'],
 'Technology': ['Copiers', 'Accessories', 'Machines', 'Phones']}
```

```python
# sales['category'] = category_subcategory_map_dataframe(df=sales, Category_dict=Category_dict, main_col='category', sub_col='sub_category')
# sales["category"]
```

### Q.3 ) Fill the missing values in sales column by mean. Return output series.

```python
sales["sales"].mean()
```

Output:
```
246.43182388324536
```

<!-- grade_id: cell-7ef02c9d67c20af2, points: 0, grade: false, solution: true -->
```python
def fill_missing(data,col):
    # your code here
    # Calculate the mean of the column
    mean_val = data[col].mean()
    
    # Fill missing values with the mean
    data[col] = data[col].fillna(mean_val)
    
    return data[col]    
```

```python
sales['sales']=fill_missing(data=sales,col='sales')
```

<!-- grade_id: cell-519682f18719b5ec, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
sales['sales'].isnull().sum()
```

Output:
```
0
```

### Q.4 ) How many orders have sales falling in the following groups 
Groups :- [0,12),[12,37),[37,113),[113,186),[186,1200),[1200,max).

**Note :-[a,b) indicates that bin includes a but excudes b.**

**Return output series**

<!-- grade_id: cell-011ab56604d6dc6d, points: 0, grade: false, solution: true -->
```python
def manual_bin(data,col):
    
    # Define the bin edges
    bins = [0, 12, 37, 113, 186, 1200, data[col].max()]
    
    # Use pandas cut function to bin the data
    data['bins'] = pd.cut(data[col], bins, right=False)
    
    # Count the number of data points in each bin
    bin_counts = data['bins'].value_counts().sort_index()
    
    return bin_counts
```

<!-- grade_id: cell-c70daab6a00fb4e5, points: 5, grade: true, solution: false -->
```python
assert type(manual_bin(data=sales,col='sales'))==pd.Series, "Please provide output in series format."
```

<!-- grade_id: cell-5553c36ce34800ac, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
manual_bin(sales,'sales') ### PASSED
```

Output:
```
[0.0, 12.0)            3708
[12.0, 37.0)           9325
[37.0, 113.0)         12761
[113.0, 186.0)         5187
[186.0, 1200.0)       18550
[1200.0, 22638.48)     1758
Name: bins, dtype: int64
```

```python
len(sales[(sales['sales'] >= 37.0) & (sales['sales'] < 113.0)])
```

Output:
```
12761
```

### Q.5 )  Sometimes, it is difficult to comprehend the sales value as a numeric value. So, the business team wants to observe the sales as Low, Medium and High categories.  Write a code to discretize the sales data and  return a series containing Low, Medium and High with the respective number of records.

<!-- grade_id: cell-8e4399b1bccc98dd, points: 0, grade: false, solution: true -->
```python
def discrete_bin(data,col):
    
    # Define the bin edges based on the quantiles
    bins = [data[col].min(), data[col].quantile(0.25), data[col].quantile(0.75), data[col].max()]
    
    # Define the bin labels
    labels = ['Low', 'Medium', 'High']
    
    # Use pandas cut function to bin the data
    data['bins'] = pd.cut(data[col], bins, labels=labels, include_lowest=True)
    
    # Count the number of data points in each bin
    bin_counts = data['bins'].value_counts().sort_index()
    
    return bin_counts
```

<!-- grade_id: cell-8f418bc782917abb, points: 5, grade: true, solution: false -->
```python
assert type(discrete_bin(data=sales,col='quantity'))==pd.Series, "Please provide output in series format."
```

<!-- grade_id: cell-4a156f92f5438b15, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
sales['quantity'].min(), sales['quantity'].quantile(0.25), sales['quantity'].quantile(0.75), sales['quantity'].max()
```

Output:
```
(1, 2.0, 5.0, 14)
```

```python
sales['quantity'].describe()
```

Output:
```
count    51290.000000
mean         3.476545
std          2.278766
min          1.000000
25%          2.000000
50%          3.000000
75%          5.000000
max         14.000000
Name: quantity, dtype: float64
```

```python
discrete_bin(sales,'quantity')  ##PASSED 50%
```

Output:
```
Low       21711
Medium    20949
High       8630
Name: bins, dtype: int64
```

### Q.6. In order to perform further analysis, we intend to utilize the 'order_priority' variable, which has categorical values ('First Class'(0) > 'Second Class'(1) > 'Standard Class'(2) > 'Same Day'(3)). To facilitate this analysis, we need to convert 'order_priority' into a numerical feature. Please encode 'order_priority' into a numerical format.Return output series.

```python
sales["ship_mode"].value_counts()
```

Output:
```
Standard Class    29842
Second Class       9985
First Class        7272
Same Day           2621
Name: ship_mode, dtype: int64
```

```python
#sales["order_priority"].value_counts()
```

<!-- grade_id: cell-d684beb2bc73d5a6, points: 0, grade: false, solution: true -->
```python
## Select Appropriate priority

priority = {'First Class': 0, 'Second Class': 1, 'Standard Class': 2, 'Same Day': 3}


def transforming_ordered_var(data,col,priority,nan_val):
    
    # Replace the categorical values with the corresponding numerical values
    data[col] = data[col].map(priority)
    
    # Fill NaN values with the specified value
    data[col].fillna(nan_val, inplace=True)
    
    return data[col]    
```

```python
sales['shipping_mode_encoded']=transforming_ordered_var(data=sales,col='ship_mode',priority=priority,nan_val=5)
```

```python
#sales['order_priority_encoded']=transforming_ordered_var(data=sales,col='order_priority', priority=priority,nan_val=5)
```

<!-- grade_id: cell-f1ba60c8543b3583, points: 20, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python
sales['shipping_mode_encoded'].value_counts()
```

Output:
```
2.0    29842
1.0     9985
0.0     7272
3.0     2621
5.0     1570
Name: shipping_mode_encoded, dtype: int64
```

```python
transforming_ordered_var(data=sales,col='ship_mode',priority=priority,nan_val=5).value_counts().loc[lambda x:x.index==5]==1570
```

Output:
```
5.0    False
Name: ship_mode, dtype: bool
```

```python
#sales['order_priority_encoded'].value_counts()
```

### Q.8) Outliers can be defined in different ways. For each of the following definitions of outliers, what percent of the values in the  'sales' are considered outliers? Give the answer as a percent rounded upto two decimal places.

Note -
1. 'IQR_detection' : Less than (Q1 - 1.5*IQR) and greater than (Q3 + 1.5*IQR) are considered outliers

2. 'percentile_detection': Below the 3rd percentile and above the 97th percentile are considered outliers

3. 'mean_SD' : Less than (mean - 3*SD) and greater than (mean + 3*SD) are considered outliers 

<!-- grade_id: cell-c1830537e6d495e9, points: 0, grade: false, solution: true -->
```python
def detect_outliers(data,col,method):
    if method=='IQR_detection':
        # Calculate the first quartile (Q1) and the third quartile (Q3)
        Q1 = data[col].quantile(0.25)
        Q3 = data[col].quantile(0.75)
        
        # Calculate the interquartile range (IQR)
        IQR = Q3 - Q1
        
        # Determine the lower and upper bounds for outliers
        lower_bound = Q1 - 1.5 * IQR
        upper_bound = Q3 + 1.5 * IQR
        
        # Count the number of outliers
        outliers = (data[col] < lower_bound) | (data[col] > upper_bound)
    elif method=='percentile_detection':
        # Calculate the 3rd and 97th percentiles
        percentile_3 = data[col].quantile(0.03)
        percentile_97 = data[col].quantile(0.97)
        
        # Determine the lower and upper bounds for outliers
        lower_bound = percentile_3
        upper_bound = percentile_97
        
        # Count the number of outliers
        outliers = (data[col] < lower_bound) | (data[col] > upper_bound)
    elif method=='mean_SD':
        # Calculate the mean and standard deviation (SD)
        mean = data[col].mean()
        std_dev = data[col].std()
        
        # Determine the lower and upper bounds for outliers
        lower_bound = mean - 3 * std_dev
        upper_bound = mean + 3 * std_dev
        
        # Count the number of outliers
        outliers = (data[col] < lower_bound) | (data[col] > upper_bound)

    outliers_perc = float((outliers.sum() / len(data[col])) * 100)
    
    outliers_perc = round(outliers_perc,2)
    
    return outliers_perc
```

<!-- grade_id: cell-618cb2f39f3bdbb8, points: 5, grade: true, solution: false -->
```python
assert type(detect_outliers(data=sales,col='sales',method='IQR_detection'))==np.float, 'Make sure that output is returned as float'
```

Output:
```
Cell In[43], line 1
----> 1 assert type(detect_outliers(data=sales,col='sales',method='IQR_detection'))==np.float, 'Make sure that output is returned as float'

File C:\ProgramData\miniconda3\lib\site-packages\numpy\__init__.py:305, in __getattr__(attr)
    300     warnings.warn(
    301         f"In the future `np.{attr}` will be defined as the "
    302         "corresponding NumPy scalar.", FutureWarning, stacklevel=2)
    304 if attr in __former_attrs__:
--> 305     raise AttributeError(__former_attrs__[attr])
    307 # Importing Tester requires importing all of UnitTest which is not a
    308 # cheap import Since it is mainly used in test suits, we lazy import it
    309 # here to save on the order of 10 ms of import time for most users
    310 #
    311 # The previous way Tester was imported also had a side effect of adding
    312 # the full `numpy.testing` namespace
    313 if attr == 'testing':

AttributeError: module 'numpy' has no attribute 'float'.
`np.float` was a deprecat
```

<!-- grade_id: cell-bfd2679b1e2a3929, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.10) Detect upper end outliers in profit column by 'IQR_detection'. Return number of upper end outliers.

<!-- grade_id: cell-96b2becd9fec86dd, points: 0, grade: false, solution: true -->
```python
def detect_upperend_outliers(data,col):
    
    # Calculate the first quartile (Q1) and the third quartile (Q3)
    Q1 = data[col].quantile(0.25)
    Q3 = data[col].quantile(0.75)
        
    # Calculate the interquartile range (IQR)
    IQR = Q3 - Q1
        
    # Determine the lower and upper bounds for outliers
    lower_bound = Q1 - 1.5 * IQR
    upper_bound = Q3 + 1.5 * IQR
        
    # Count the number of outliers
    upper_outliers = int((data[col] > upper_bound).sum())
        
    return upper_outliers
```

<!-- grade_id: cell-9623402867c6a9ef, points: 5, grade: true, solution: false -->
```python
assert type(detect_upperend_outliers(data=sales,col='profit'))==int , 'No of outliers should be an integer , make sure that you are returning output in integer format.'
```

<!-- grade_id: cell-5e43e81906b9a713, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.11) Floor & Capping outliers at 97% at 3%. Return the processed column as output.

<!-- grade_id: cell-f5527eddcb6c9757, points: 0, grade: false, solution: true -->
```python
def cap_and_floor_column(column, cap_percentile, floor_percentile):
    # your code here
    return capped_and_floored_column ## Capped ,Floored series
```

```python
sales['profit_outlier_capped'] = cap_and_floor_column(column=sales['profit'], cap_percentile=97, floor_percentile=3)
```

Output:
```
Cell In[49], line 1
----> 1 sales['profit_outlier_capped'] = cap_and_floor_column(column=sales['profit'], cap_percentile=97, floor_percentile=3)

Cell In[48], line 3, in cap_and_floor_column(column, cap_percentile, floor_percentile)
      1 def cap_and_floor_column(column, cap_percentile, floor_percentile):
      2     # your code here
----> 3     return capped_and_floored_column

NameError: name 'capped_and_floored_column' is not defined
```

<!-- grade_id: cell-e8e8dbefcd633180, points: 5, grade: true, solution: false -->
```python
assert round(np.quantile(sales['profit_outlier_capped'],0.97),4)==301.9693, 'Make sure that you are flooring and capping with correct percentiles & returning the correct output.'
assert round(np.quantile(sales['profit_outlier_capped'],0.03),4)==-144.632, 'Make sure that you are flooring and capping with correct percentiles & returning the correct output.'
```

Output:
```
File C:\ProgramData\miniconda3\lib\site-packages\pandas\core\frame.py:3807, in DataFrame.__getitem__(self, key)
   3805 if self.columns.nlevels > 1:
   3806     return self._getitem_multilevel(key)
-> 3807 indexer = self.columns.get_loc(key)
   3808 if is_integer(indexer):
   3809     indexer = [indexer]

File C:\ProgramData\miniconda3\lib\site-packages\pandas\core\indexes\base.py:3804, in Index.get_loc(self, key, method, tolerance)
   3802     return self._engine.get_loc(casted_key)
   3803 except KeyError as err:
-> 3804     raise KeyError(key) from err
   3805 except TypeError:
   3806     # If we have a listlike key, _check_indexing_error will raise
   3807     #  InvalidIndexError. Otherwise we fall through and re-raise
   3808     #  the TypeError.
   3809     self._check_indexing_error(key)

KeyError: 'profit_outlier_capped'
```

<!-- grade_id: cell-d317f90221447667, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

### Q.12) Business team is interested in knowing how many vendors have profit_outlier_capped >10000 & vendors with profit_outlier_capped <10000. Return the dataframe with 3 columns vendor  , profit_outlier_capped , profit_bins & sort it with 'profit_outlier_capped' in descending order.

Instructions:-
1. Bins should be in (a,b] format. [a,b) indicates that bin includes a but excudes b])
2. Make sure that for 1st interval it is **left inclusive** i.e [a,b].
3. Segment the ''profit_outlier_capped'' variable in 2 bins. (first bin with profit>10000 ,second bin with profit<10000).

<!-- grade_id: cell-77f6b02474e37c07, points: 0, grade: false, solution: true -->
```python
def profit_vendor(data,col):
    # your code here
```

Output:
```
Cell In[52], line 2
    # your code here
                    ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: cell-e3d429861a4da377, points: 5, grade: true, solution: false -->
```python
assert type(profit_vendor(data=sales,col='profit_outlier_capped'))==pd.DataFrame ,'Make sure that your output is a dataframe '
assert profit_vendor(data=sales,col='profit_outlier_capped').shape==(65,3) ,'Make sure that shape of your dataframe is correct'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[53], line 1
----> 1 assert type(profit_vendor(data=sales,col='profit_outlier_capped'))==pd.DataFrame ,'Make sure that your output is a dataframe '
      2 assert profit_vendor(data=sales,col='profit_outlier_capped').shape==(65,3) ,'Make sure that shape of your dataframe is correct'

NameError: name 'profit_vendor' is not defined
```

<!-- grade_id: cell-05010ee8edfc15a7, points: 5, grade: true, solution: false -->
```python
# autograder cells , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```

```python

```