### Graded Lab 5

Hello ! Welcome to Graded Lab of Module 5.

In the last assignment we had worked on missing values & outliers.
Its high time when we look at the feature engineering & Principal component analysis.

**So lets go ahead with the last graded assignment for the course !**

In case if you are not able to recollect the problem description and data description then mentioning it below.

Lets look at the problem statement,

*Client: ABC Retail, Incorporated, rest-of-the-world division* 

***Project name: Online retail sales analysis*** 

An online retailer, ABC, Inc., operates in nearly 100 countries worldwide, selling furniture, office supplies and technology products to customers in three segments: consumer, corporate and home office. ABC, Inc. is a US-based company, and it has two major divisions: US and rest of the world. We are working with the rest of the world division of the company. 

They have provided us with online sales transaction data from 2011 to 2014.

We are given 3 datasets:-

1. Data on each sale; 51290 records; all data in US dollars
It contains fields like
**order_id** (identifier) ,order_date ,ship_date ,ship_mode ,**customer_id**(identifier) ,product_id ,category ,sub_category ,product_name ,sales ,quantity ,discount ,profit ,shipping_cost ,order_priority ,**vendor_code** (identifier) 


2. Data on the customers; 1590 records 
It contains fields like
**customer_id** (identifier) ,customer_name ,city ,state ,country ,postal_code ,segment ,market ,region 

3. Data on vendors who supply the retailer; 65 records 
It contains fields like
vendor ,**vendor_code** (identifier) 

We need to analyze the data and need to provide answer to different questions asked by company officials.

# Instructor update: Jan 29, 2024

```python
# importing libraries
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
from scipy.stats import norm
from itertools import combinations
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA

### Reading sales data
sales = pd.read_csv('sales_data.csv')

### Reading customer data
cust = pd.read_csv(r'customers.csv',encoding='iso-8859-1')

### Reading vendor data
vend = pd.read_csv(r'vendors.csv')
sales.head()
```

In order to solve the next questions , we need to combine all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. So here we have written a data processing function.
There are 2 tasks which are to be performed.
1. Merge/ Join all the 3 datasets into a single dataframe such that every details of sales dataframe are intact. (Understand which should be the joining key , type of join , refer .merge() function of pandas)
2. Convert 'order_date' into a datetime column.
**Return output as a dataframe**

<!-- grade_id: cell-171f8d30f5d84056, points: 0, grade: false, solution: true -->
```python
#### data_merging & order_date processing , data1 will be sales , data2 will be customer dataset & data3 will be vendor dataset.

def data_process(data1,data2,data3):
    
    return data
```

```python
sales= data_process(data1=sales.copy(),data2=cust.copy(),data3=vend.copy())
```

<!-- grade_id: cell-61ebe95bf598597b, points: 0, grade: true, solution: false -->
```python
assert sales['order_date'].dtypes=='<M8[ns]' ,'Make sure if you have converted order_date into a datetime format correctly or not.'
assert sales.shape== (51290,26) ,'Checking size and shape of dataframe after merging is a very important check.'
```

### Q.1) Create a feature for identifying which orders are ordered on weekends. Flag will take value 1 if order is ordered on weekends & 0 otherwise. (Sat & Sun are weekends). Return  the output series.

Hint: Think of using weekday() function, consider'order_date' as the required columns for computations. 

<!-- grade_id: cell-da758c3c17ca0ef3, points: 0, grade: false, solution: true -->
```python
def weekend_flag(data):
    # your code here
```

```python
sales['weekend_orders']=weekend_flag(data=sales)
```

<!-- grade_id: cell-12c0ef94c14c3ef0, points: 7, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert all(sales['weekend_orders'].unique()==np.array([0, 1])) ,'Make sure that you have created flag correctly'

```

### Q.2) We wish to create a feature called 'discount_per_quantity'. Write a functional code to compute ratio of 2 columns. Return  the output series.

<!-- grade_id: cell-bf222819379099c4, points: 0, grade: false, solution: true -->
```python
def feature_1(data,col1,col2):
    # your code here
```

```python
sales['discount_per_quantity']= feature_1(data=sales,col1='discount',col2='quantity')
```

<!-- grade_id: cell-1d1c9c0070ebdf4e, points: 7, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert round(sales['discount_per_quantity'].mean(),4)==0.0643,'Make sure that you have created feature correctly.'

```

### Q.3) Create a feature for finding cumulative sales for every customer. Return  the output series.

<!-- grade_id: cell-23b8a0fddaa72e30, points: 0, grade: false, solution: true -->
```python
def cumulative_sales(data):
    # your code here
```

```python
### Sorting the dataframe wrto customer_id & order_date
sales.sort_values(['customer_id','order_date'],ascending=[1,1],inplace=True)
### Applying cumulative_sales function
sales['cumulative_sales']=cumulative_sales(data=sales)
```

<!-- grade_id: cell-f88f29cf32252ab7, points: 10, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert round(sales[sales['customer_id']=='ZC-21910']['cumulative_sales'].max(),2)==28472.82
```

### Q.4) Please demonstrate how to enhance a given dataset by introducing a new feature that records the frequency of same-day purchases made by each customer. The desired outcome should be a modified dataframe containing a variable named 'same_day_purchase_frequency' which accurately portrays the same-day purchase behavior of individual customers.


Hint:

1.Same_day_purchase_frequency will be the frequency of the orders made on same day. If customer 'ABC' has placed 2 orders on 2nd Jan 23 then 'same_day_purchase_frequency'for this customer for 2nd Jan 23 will be 2. This is still at Customer * Date level

2.If customer  'ABC' has made say 10 such frequency transactions where frequency of same day purchase >1. Then'same_day_frequency_purchase' for customer will be 10. This is our final answer at Customer level.  

<!-- grade_id: cell-11502c58370a30cb, points: 0, grade: false, solution: true -->
```python
def same_day_purchase(data):
    # your code here
    
    # Group by 'customer_id' and date, and count the orders, store count of orders in variable named 'order_count'
    order_counts = None
    
    # Filter for same-day purchases (order_count > 1)
    same_day_purchases = order_counts[order_counts['order_count'] > None]
    
    
    # Group by 'customer_id' and calculate the same-day purchase frequency , name that column as 'same_day_purchase_frequency'
    same_day_purchase_frequency = None
    
    # Merge the result with the original DataFrame to add the same-day purchase frequency as a new column  
    data = None


    # Fill missing values with 0 (for customers with no same-day purchases)
    None

    return data
```

<!-- grade_id: cell-1a74a24a4da6dd2d, points: 30, grade: true, solution: false -->
```python
sales_new=same_day_purchase(data=sales.copy())
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert float(sales_new[sales_new['customer_id']=='ZC-11910']['same_day_purchase_frequency'].unique())==0,'Make sure that you have created variable correctly.'

```

### Q.5) Create a variable which will calculate weighted sales per market. Return  the output series.
Weights for every market are :-'LATAM':1, 'EMEA':1.3, 'Africa':2.7, 'Canada':1.4, 'EU':2.8, 'APAC':3.5

<!-- grade_id: cell-47084e349b64bd0d, points: 0, grade: false, solution: true -->
```python
def weighted_sales(data):
    # Write a lambda function & apply it to dataframe.Kindly ignore raiseNotImplementError
    # your code here
```

```python
sales['weighted_market_sales']=weighted_sales(data=sales)
```

<!-- grade_id: cell-389d62fab17c8c2f, points: 10, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert round(float(sales[(sales['market']=='Africa') & (round(sales['sales'],4)==58.83)]['weighted_market_sales'].unique()),2)==158.84, 'Make sure that ypu have created column correctly.'
```

Most of the variables in our data are categorical in nature & for further anlysis purpose if we convert them by ordinal encoding or one hot encoding, we will be left with too many variables.
If we have too many variables in the data , it might cause trouble in later stages like machine learning model building.
1. Having too many variables can lead to overfitting and poor generalization to new data.
2. It can make it difficult to analyze and interpret the data and identify the most important variables.
3. It can increase computational complexity and time required to build and train a model.

How to deal with this ?

Wait , we have learned about **Principal component analysis** , we can leverage that method here for **dimensionality reduction.**

### Q.6 ) Complete the following code for PCA implementation. Return the no of PC's required for capturing the 80% of variation, variance explained by the PC's .

<!-- grade_id: cell-033f0d1768c98329, points: 0, grade: false, solution: true -->
```python
def PCA(data,target=0.80):
    # your code here
    
    # Treating ordinal variable ,
    # order_priority : 'Medium':2, 'Low':1, 'High':3, 'Critical':4 
    data['order_priority_encoded']=data['order_priority'].map(None) # Write your code in place of None

    
    # Treating missing values
    # Impute the missing values in 'order_prioirty_encoded','ship_mode' by mode, 
    data['order_priority_encoded']=None # Write your code in place of None
    data['ship_mode']=None # Write your code in place of None

    
    # Treating Nominal variables ,'ship_mode','category','sub_category' (while encoding them make sure to select k-1 dummy variables for a variable with k categories)
    data=None # Write your code in place of None

    # Final data selection
    # Select all the variables with 'sales','quantity','discount','profit','shipping_cost','order_priority_encoded' & dummies in 'ship_mode','category','sub_category'.
    data_1=data[['sales','quantity','discount','profit','shipping_cost','order_priority_encoded','ship_mode_Same Day','ship_mode_Second Class','ship_mode_Standard Class','category_Office Supplies','category_Technology','sub_category_Appliances','sub_category_Art','sub_category_Binders','sub_category_Bookcases','sub_category_Chairs','sub_category_Copiers','sub_category_Envelopes','sub_category_Fasteners','sub_category_Furnishings','sub_category_Labels','sub_category_Machines','sub_category_Paper','sub_category_Phones','sub_category_Storage','sub_category_Supplies','sub_category_Tables']]

    
    # Standardize the data
    sc = StandardScaler()
    
    cols = ['sales','quantity','discount','profit','shipping_cost','order_priority_encoded']
    for col in cols:
        data_1[col] = None # Write your code in place of None

    
    from sklearn.decomposition import PCA
    
    # Create a PCA instance    
    pca = PCA()
    
    # Fit PCA to the scaled data
    pca.fit(None) # Write your code in place of None
    
    #  Calculate cumulative explained variance
    explained_variance_ratio = None # Write your code in place of None
    cumulative_variance = None # Write your code in place of None
    
    
    target_variation=target
    n_components = np.argmax(cumulative_variance >= target_variation) + 1
 
    
    # Excellent we got no of pcs which are required for capturing the target variation in the data
    
    ### Fitting PCA again with no of pc's we got above.
    pca = PCA(n_components=n_components)
    
    # Fit PCA to the scaled data
    pca.fit(None) # Write your code in place of None
    
    #  Calculate cumulative explained variance
    explained_variance_ratio = None # Write your code in place of None
       
    return n_components,explained_variance_ratio 
```

<!-- grade_id: cell-095ce6ebc95af61e, points: 30, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
assert type(PCA(data=sales.copy())[0])==np.int64 , 'Number of components required for capturing target variation should be in integer format.'
assert type(PCA(data=sales.copy())[1])==np.ndarray , 'Type of variation captured should be in an array format.'
assert [round(x,4) for x in list(PCA(data=sales.copy())[1])][0]==0.3012, 'Make sure that you are returning correct values for output.'

```

### Q.7) How much variation was explained by 3rd principal component ? 
**Output should be rounded upto 4 digits**

<!-- grade_id: cell-d668bca10bcf9149, points: 0, grade: false, solution: true -->
```python
def var_expl(variation_explained):
    # Kindly ignore raiseNotImplementError   
    # your code here
```

<!-- grade_id: cell-cf2fbc4c16e8ba3d, points: 6, grade: true, solution: false -->
```python
# autograder cell , please do not alter/ delete /edit this cell,Kindly ignore this cell.
```