# PROBLEM STATEMENT 

Anova Insurance, a global health insurance company, seeks to optimize its insurance policy premium pricing based on the health status of applicants. Understanding an applicant's health condition is crucial for two key decisions:
- Determining eligibility for health insurance coverage.
- Deciding on premium rates, particularly if the applicant's health indicates higher risks.

Your objective is to Develop a predictive model that utilizes health data to classify individuals as 'healthy' or 'unhealthy'. This classification will assist in making informed decisions about insurance policy premium pricing.

# OVERVIEW 

The dataset contains 10,000 rows and 20 columns (original data without preprocessing), the no. of columns becomes 23 post preprocessing because of encoding, the 23 columns includes both numerical and categorical variables. Here is the data dictionary.

- Age: Represents the age of the individual. Negative values seem to be present, which might indicate data entry errors or a specific encoding used for certain age groups.

- BMI (Body Mass Index): A measure of body fat based on height and weight. Typically, a BMI between 18.5 and 24.9 is considered normal.

- Blood_Pressure: Represents systolic blood pressure. Normal blood pressure is usually around 120/80 mmHg.

- Cholesterol: This is the cholesterol level in mg/dL. Desirable levels are usually below 200 mg/dL.

- Glucose_Level: Indicates blood glucose levels. It might be fasting glucose levels, with normal levels usually ranging from 70 to 99 mg/dL.

- Heart_Rate: The number of heartbeats per minute. Normal resting heart rate for adults ranges from 60 to 100 beats per minute.

- Sleep_Hours: The average number of hours the individual sleeps per day.

- Exercise_Hours: The average number of hours the individual exercises per day. 

- Water_Intake: The average daily water intake in liters.

- Stress_Level: A numerical representation of stress level.

- Target: This is a binary outcome variable, with '1' indicating 'Unhealthy' and '0' indicating 'Healthy'.

- Smoking: A categorical variable indicating smoking status. Contains values - (0,1,2) which specify the regularity of smoking with 0 being no smoking and 2 being regular smmoking.

- Alcohol: A categorical variable indicating alcohol consumption status. Contains values - (0,1,2) which specify the regularity of alcohol consumption with 0 being no consumption quality and 2 being regular consumption.

- Diet: A categorical variable indcating the quality of dietary habits. Contains values - (0,1,2) which specify the quality of the habit with 0 being poor diet quality and 2 being good quality.

- MentalHealth: Possibly a measure of mental health status. Contains values - (0,1,2) which specify the severity of the mental health with 0 being fine and 2 being highly severe

- PhysicalActivity: A categorical variable indicating levels of physical activity. Contains values - (0,1,2) which specify the instensity of the medical history with 0 being no Physical Activity and 2 being regularly active.

- MedicalHistory: Indicates the presence of medical conditions or history. Contains values - (0,1,2) which specify the severity of the medical history with 0 being nothing and 2 being highly severe.

- Allergies: A categorical variable indicating allergy status. Contains values - (0,1,2) which specify the severity of the allergies with 0 being nothing and 2 being highly severe.

- Diet_Type: Categorical variable indicating the type of diet an individual follows. Contains values(Vegetarian, Non-Vegetarian, Vegan). 
- (this column has been encoded into three different columns during the preprocessing stage)
 - Diet_Type_Vegan,Diet_Type_Vegetarian

- Blood_Group: Indicates the blood group of the individual Contains values (A, B, AB, O), this column values are encoded too .

It is clear from the above description that the predictor variable is the 'Target' column.


# -----------------------------------------------------------------------------

## Guidelines to follow in this notebook 
- The name of the dataframe should be df 
- Keep the seed value 42
- Names of training and testing variables should be X_train, X_test, y_train, y_test
- Keep the name of model instance as "model", e.g. model = DecisionTreeClassifer()
- Keep the predictions on training and testing data in a variable named y_train_pred and y_test_pred respectively.

# -------------------------------------------------------------------------------


Let us begin with importing the necessary libraries.

```python

#import relevant data libraries 

import pandas as pd
pd.set_option('display.max_rows', 500)
pd.set_option('display.max_columns', 500)
pd.set_option('display.width', 1000)

from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
from sklearn.ensemble import BaggingClassifier

from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import f1_score

from collections import OrderedDict
from sklearn.datasets import make_classification
from sklearn.ensemble import RandomForestClassifier
import matplotlib.pyplot as plt
from sklearn.ensemble import AdaBoostClassifier
from sklearn.ensemble import GradientBoostingClassifier
!pip install xgboost
import xgboost as xgb
!pip install lightgbm
from lightgbm import LGBMClassifier
!pip install catboost 
from catboost import CatBoostClassifier

```

Output:
```
Requirement already satisfied: xgboost in c:\programdata\miniconda3\lib\site-packages (2.0.3)
Requirement already satisfied: numpy in c:\programdata\miniconda3\lib\site-packages (from xgboost) (1.24.4)
Requirement already satisfied: scipy in c:\programdata\miniconda3\lib\site-packages (from xgboost) (1.10.1)
Requirement already satisfied: lightgbm in c:\programdata\miniconda3\lib\site-packages (4.2.0)
Requirement already satisfied: numpy in c:\programdata\miniconda3\lib\site-packages (from lightgbm) (1.24.4)
Requirement already satisfied: scipy in c:\programdata\miniconda3\lib\site-packages (from lightgbm) (1.10.1)
Requirement already satisfied: catboost in c:\programdata\miniconda3\lib\site-packages (1.2.2)
Requirement already satisfied: graphviz in c:\programdata\miniconda3\lib\site-packages (from catboost) (0.20.1)
Requirement already satisfied: matplotlib in c:\programdata\miniconda3\lib\site-packages (from catboost) (3.6.0)
Requirement already satisfied: numpy>=1.16.0 in c:\progra
```

## Load the dataset

```python
#Load the dataset
df = pd.read_csv('Healthcare_Dataset_Preprocessednew.csv')
df.head()
```

Output:
```
Age   BMI  Blood_Pressure  Cholesterol  Glucose_Level  Heart_Rate  Sleep_Hours  Exercise_Hours  Water_Intake  Stress_Level  Target  Smoking  Alcohol  Diet  MentalHealth  PhysicalActivity  MedicalHistory  Allergies  Diet_Type__Vegan  Diet_Type__Vegetarian  Blood_Group_AB  Blood_Group_B  Blood_Group_O
0   2.0  26.0           111.0        198.0           99.0        72.0          4.0             1.0           5.0           5.0       1        2        2     1             2                 1               0          1             False                   True            True          False          False
1   8.0  24.0           121.0        199.0          103.0        75.0          2.0             1.0           2.0           9.0       1        0        1     1             2                 1               2          2             False                  False            True          False          False
2  81.0  27.0           147.0        203.0          100.0        74.0         10.0       
```

```python
#shape of data
df.shape

```

Output:
```
(9549, 23)
```

```python
# Column names in the dataset
df.columns
```

Output:
```
Index(['Age', 'BMI', 'Blood_Pressure', 'Cholesterol', 'Glucose_Level', 'Heart_Rate', 'Sleep_Hours', 'Exercise_Hours', 'Water_Intake', 'Stress_Level', 'Target', 'Smoking', 'Alcohol', 'Diet', 'MentalHealth', 'PhysicalActivity', 'MedicalHistory', 'Allergies', 'Diet_Type__Vegan', 'Diet_Type__Vegetarian', 'Blood_Group_AB', 'Blood_Group_B', 'Blood_Group_O'], dtype='object')
```

# Separate the indpendent features in the dataframe 'X' and  the target in a variable 'y '


<!-- grade_id: cell-03aeb917f862eb43, points: 0, grade: false, solution: true -->
```python
# your code here
X = df.drop(['Target'], axis=1)
y = df['Target']
```

```python
X.shape , y.shape
```

Output:
```
((9549, 22), (9549,))
```

<!-- grade_id: cell-c1a78a0af3eaf3d3, points: 2, grade: true, solution: false -->
```python

```

# Splitting Dataset into Train and Test Sets


<!-- grade_id: cell-2f0f4b155b1a7859, points: 0, grade: false, solution: true -->
```python
# Splitting the dataset into training and testing sets keeping the test size as 25% and seed value as 42

# your code here
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25, random_state=42)
```

<!-- grade_id: cell-a00e8554b95f924a, points: 2, grade: true, solution: false -->
```python

```

# Create a simple Decision tree 

The first step is to train the model using simple decision tree 

<!-- grade_id: cell-7f0afa6c40ec7eb2, points: 0, grade: false, solution: true -->
```python
# Initialize a simple Decision Tree classifier with depth 15 and seed 42. Name it 'model' and then fit it 
# your code here
model = DecisionTreeClassifier(max_depth=15, random_state=42)
```

<!-- grade_id: cell-af006460199d3a77, points: 3, grade: true, solution: false -->
```python
## Begin hidden test
assert model.max_depth == 15, "Max_depth is not set to 15"
## End hidden test
```

<!-- grade_id: cell-94f0a1add824afce, points: 0, grade: false, solution: true -->
```python
# your code here
model.fit(X_train, y_train.values.ravel())
```

Output:
```
DecisionTreeClassifier(max_depth=15, random_state=42)
```

```python
y_train_pred = model.predict(X_train)
y_test_pred = model.predict(X_test)
```

# Evaluate model performance

After creating model and getting the predictions on the test set, 
 calculate f1 score for evaluating the performance of the model 

<!-- grade_id: cell-6e095f8c130aff64, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import f1_score

# your code here
f1_score(y_test,y_test_pred)
```

Output:
```
0.903097696584591
```

<!-- grade_id: cell-07f02dd4841a52e7, points: 1, grade: true, solution: false -->
```python

```

# APPLY ADABOOST ALGORITHM FOR TRAINING 

After creating simple decision tree, its now time to create classifier model using AdaBoostClassifier 

<!-- grade_id: cell-d9be0970df96e6c5, points: 0, grade: false, solution: true -->
```python

# your code here
model = AdaBoostClassifier(random_state=42)
# Train the AdaBoost model
model.fit(X_train, y_train.values.ravel())
```

Output:
```
AdaBoostClassifier(random_state=42)
```

After training the model using AdaBoostClassifier , do prediction on the test data and calculate f1 score for evaluation

```python
# Predict on training and validation data
y_pred_train = model.predict(X_train)
y_pred_test = model.predict(X_test)
```

<!-- grade_id: cell-ffa36d2556fb0641, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import f1_score
# your code here
f1_score(y_test,y_test_pred)
```

Output:
```
0.903097696584591
```

<!-- grade_id: cell-2235d7485226df2d, points: 2, grade: true, solution: false -->
```python

```

# APPLY GRADIENT BOOSTING 

### Now use GradientBoostingClassifier for training the model 


<!-- grade_id: cell-41421b7013a91946, points: 0, grade: false, solution: true -->
```python
model = GradientBoostingClassifier(random_state=42)
model.fit(X_train, y_train.values.ravel())
```

Output:
```
GradientBoostingClassifier(random_state=42)
```

## Evaluation

Prediciton on training and testing data using GradientBoostingClassifier and then calculate f1 score 

<!-- grade_id: cell-b7277ccd9508aa4e, points: 0, grade: false, solution: true -->
```python
y_pred_train = model.predict(X_train)
y_pred_test = model.predict(X_test)
```

<!-- grade_id: cell-9ea03aa00c137df1, points: 0, grade: false, solution: true -->
```python
f1_score(y_test,y_test_pred)
```

Output:
```
0.903097696584591
```

<!-- grade_id: cell-83b63df4d495b4db, points: 4, grade: true, solution: false -->
```python

```

# APPLY XGBOOST 

### Here, use XGBClassifier to train the model.


<!-- grade_id: cell-8fa1a789e3a71b45, points: 0, grade: false, solution: true -->
```python
from xgboost import XGBClassifier

model = XGBClassifier(random_state=42)
model.fit(X_train, y_train.values.ravel())
```

Output:
```
XGBClassifier(base_score=None, booster=None, callbacks=None,
              colsample_bylevel=None, colsample_bynode=None,
              colsample_bytree=None, device=None, early_stopping_rounds=None,
              enable_categorical=False, eval_metric=None, feature_types=None,
              gamma=None, grow_policy=None, importance_type=None,
              interaction_constraints=None, learning_rate=None, max_bin=None,
              max_cat_threshold=None, max_cat_to_onehot=None,
              max_delta_step=None, max_depth=None, max_leaves=None,
              min_child_weight=None, missing=nan, monotone_constraints=None,
              multi_strategy=None, n_estimators=None, n_jobs=None,
              num_parallel_tree=None, random_state=42, ...)
```

<!-- grade_id: cell-232929f197a03beb, points: 0, grade: false, solution: true -->
```python
y_pred_train = model.predict(X_train)
y_pred_test = model.predict(X_test)
```

<!-- grade_id: cell-acce8aa2748578d5, points: 0, grade: false, solution: true -->
```python
f1_score(y_test,y_test_pred)
```

Output:
```
0.903097696584591
```

<!-- grade_id: cell-d66e72a32f15e595, points: 4, grade: true, solution: false -->
```python

```

# APPLY LIGHTGBM 

### TRAIN MODEL WITH LGBMClassifier this time 

<!-- grade_id: cell-c5ce26b288580bce, points: 0, grade: false, solution: true -->
```python
from lightgbm import LGBMClassifier

model = LGBMClassifier(random_state=42)
model.fit(X_train, y_train.values.ravel())
```

Output:
```
[LightGBM] [Info] Number of positive: 3709, number of negative: 3452
[LightGBM] [Info] Auto-choosing row-wise multi-threading, the overhead of testing was 0.000266 seconds.
You can set `force_row_wise=true` to remove the overhead.
And if memory is not enough, you can set `force_col_wise=true`.
[LightGBM] [Info] Total Bins 396
[LightGBM] [Info] Number of data points in the train set: 7161, number of used features: 22
[LightGBM] [Info] [binary:BoostFromScore]: pavg=0.517944 -> initscore=0.071809
[LightGBM] [Info] Start training from score 0.071809
LGBMClassifier(random_state=42)
```

```python
# Predict on the training and testing data
y_train_pred = model.predict(X_train)
y_test_pred = model.predict(X_test)

```

<!-- grade_id: cell-fb45246d751b99c0, points: 0, grade: false, solution: true -->
```python
f1_score(y_test,y_test_pred)
```

Output:
```
0.9511811023622048
```

<!-- grade_id: cell-d7fbcddebc3cde64, points: 2, grade: true, solution: false -->
```python

```

# APPLY CATBOOST 

### Train model with CatBoostClassifier 

<!-- grade_id: cell-7f2fc5358a8b3e80, points: 0, grade: false, solution: true -->
```python
from catboost import CatBoostClassifier

model = CatBoostClassifier(random_state=42)

model.fit(X_train, y_train.values.ravel())
```

Output:
```
Learning rate set to 0.023879
0:	learn: 0.6749614	total: 156ms	remaining: 2m 35s
1:	learn: 0.6560950	total: 168ms	remaining: 1m 24s
2:	learn: 0.6380358	total: 180ms	remaining: 59.8s
3:	learn: 0.6204795	total: 185ms	remaining: 46s
4:	learn: 0.6054240	total: 190ms	remaining: 37.9s
5:	learn: 0.5910817	total: 195ms	remaining: 32.3s
6:	learn: 0.5759772	total: 199ms	remaining: 28.2s
7:	learn: 0.5627926	total: 202ms	remaining: 25s
8:	learn: 0.5511314	total: 206ms	remaining: 22.7s
9:	learn: 0.5395501	total: 210ms	remaining: 20.8s
10:	learn: 0.5284412	total: 215ms	remaining: 19.3s
11:	learn: 0.5172321	total: 219ms	remaining: 18.1s
12:	learn: 0.5091017	total: 225ms	remaining: 17.1s
13:	learn: 0.5001794	total: 228ms	remaining: 16.1s
14:	learn: 0.4927348	total: 232ms	remaining: 15.3s
15:	learn: 0.4844060	total: 237ms	remaining: 14.6s
16:	learn: 0.4759033	total: 241ms	remaining: 14s
17:	learn: 0.4676664	total: 245ms	remaining: 13.4s
18:	learn: 0.4600816	total: 249ms	remaining: 12.9s
19:	learn: 0.45
```

<!-- grade_id: cell-91aded2c00c87230, points: 0, grade: false, solution: true -->
```python
y_train_pred = model.predict(X_train)
y_test_pred = model.predict(X_test)
```

<!-- grade_id: cell-20064a5991d16a83, points: 0, grade: false, solution: true -->
```python
f1_score(y_test,y_test_pred)
```

Output:
```
0.9569680221081721
```

<!-- grade_id: cell-8b4497c5fb187ba0, points: 5, grade: true, solution: false -->
```python

```

The performance of various boosting techniques depends on multiple hyperparamters and dataset provided for training. 
With gradient boosting, good precision can be achieved in reducing errors. There are multiple advantages in using gradient boosting as compared to adaboost. Also XgBoost is again evolved version of gradient boosting. Being versatile and capability of handling large datasets makes it more popular. While CatBoost is popular to handle categorical variables more vigourously and doesnt require encoding for the same.


